A Novel 3D Visualization Method in Mice Identifies the Periportal Lamellar Complex (PLC) as a Key Regulator of Hepatic Ductal and Neuronal Branching Morphogenesis

Author response:

The following is the authors’ response to the original reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

In this manuscript, Chengjian Zhao et al. focused on the interactions between vascular, biliary, and neural networks in the liver microenvironment, addressing the critical bottleneck that the lack of high-resolution 3D visualization has hindered understanding of these interactions in liver disease.

Strengths:

This study developed a high-resolution multiplex 3D imaging method that integrates multicolor metallic compound nanoparticle (MCNP) perfusion with optimized CUBIC tissue clearing. This method enables the simultaneous 3D visualization of spatial networks of the portal vein, hepatic artery, bile ducts, and central vein in the mouse liver. The authors reported a perivascular structure termed the Periportal Lamellar Complex (PLC), which is identified along the portal vein axis. This study clarifies that the PLC comprises CD34⁺Sca-1⁺ dual-positive endothelial cells with a distinct gene expression profile, and reveals its colocalization with terminal bile duct branches and sympathetic nerve fibers under physiological conditions.

Weaknesses:

This manuscript is well-written, organized, and informative. However, there are some points that need to be clarified.

(1) After MCNP-dye injection, does it remain in the blood vessels, adsorb onto the cell surface, or permeate into the cells? Does the MCNP-dye have cell selectivity?

The experimental results showed that after injection, the MCNP series nanoparticles predominantly remained within the lumens of blood vessels and bile ducts, with their tissue distribution determined by physical perfusion. No diffusion of the dye signal into the surrounding parenchymal tissue was observed, nor was there any evidence of adsorption onto the cell surface or entry into cells. The newly added Supplementary Figure S2A–H further confirmed this feature, demonstrating that the dye signals were strictly confined to the luminal space, clearly delineating the continuous course of blood vessels and the branching morphology of bile ducts. These findings strongly support the conclusion that “MCNP dyes are distributed exclusively within the luminal compartments.”

Therefore, the MCNP dyes primarily serve as intraluminal tracers within the tissue rather than as labels for specific cell types.

(2) All MCNP-dyes were injected after the mice were sacrificed, and the mice's livers were fixed with PFA. After the blood flow had ceased, how did the authors ensure that the MCNP-dyes were fully and uniformly perfused into the microcirculation of the liver?

Thank you for the reviewer’s valuable comments. Indeed, since all MCNP dyes were perfused after the mice were euthanized and blood circulation had ceased, we cannot fully ensure a homogeneous distribution of the dye within the hepatic microcirculation. The vascular labeling technique based on metallic nanoparticle dyes used in this study offers clear imaging, stable fluorescence intensity, and multiplexing advantages; however, it also has certain limitations. The main issue is that the dye distribution within the hepatic parenchyma can be affected by factors such as lobular overlap, local tissue compression, and variations in vascular pathways, resulting in regional inhomogeneity of dye perfusion. This is particularly evident in areas where multiple lobes converge or where anatomical structures are complex, leading to local dye accumulation or over-perfusion.

In our experiments, we attempted to minimize local blockage or over-perfusion by performing PBS pre-flushing and low-pressure, constant-speed perfusion. Nevertheless, localized dye accumulation or uneven distribution may still occur in lobe junctions or structurally complex regions. Such variation represents one of the methodological limitations. Overall, the dye signals in most samples remained confined to the vascular and biliary lumens, and the distribution pattern was highly reproducible.

We have addressed this issue in the Discussion section but would like to emphasize here that, although this system has clear advantages, it remains sensitive to anatomical variability in the liver—such as lobular overlap and vascular heterogeneity. At vascular junctions, local perfusion inhomogeneity or dye accumulation may occur; therefore, injection strategies and perfusion parameters should be adjusted according to liver size and vascular condition to improve reproducibility and imaging quality. It should also be noted that the results obtained using this method primarily aim to visualize the overall and fine anatomical structures of the hepatic vascular system rather than to quantitatively reflect hemodynamic processes. In the future, we plan to combine in vivo perfusion or dynamic fluid modeling to further validate the diffusion characteristics of the dyes within the hepatic microcirculation.

(3) It is advisable to present additional 3D perspective views in the article, as the current images exhibit very weak 3D effects. Furthermore, it would be better to supplement with some videos to demonstrate the 3D effects of the stained blood vessels.

Thank you for the reviewer’s valuable comments. In response to the suggestion, we have added perspective-rendered images generated from the 3D staining datasets to provide a more intuitive visualization of the spatial morphology of the hepatic vasculature. These images have been included in Figure S2A–J. In addition, we have prepared supplementary videos (available upon request) that dynamically display the three-dimensional distribution of the stained vessels, further enhancing the spatial perception and visualization of the results.

(4) In Figure 1-I, the authors used MCNP-Black to stain the central veins; however, in addition to black, there are also yellow and red stains in the image. The authors need to explain what these stains are in the legend.

Thank you for the reviewer’s constructive comment. In Figure 1I, MCNP-Black labels the central vein (black), MCNP-Yellow labels the portal vein (yellow), MCNP-Pink labels the hepatic artery (pink), and MCNP-Green labels the bile duct (green). We have revised the Figure 1 legend to include detailed descriptions of the color signals and their corresponding structures to avoid any potential confusion.

(5) There is a typo in the title of Figure 4F; it should be "stem cell".

Thank you for the reviewer’s careful correction. We have corrected the spelling error in the title of Figure 4F to “stem cell” and updated it in the revised manuscript.

(6) Nuclear staining is necessary in immunofluorescence staining, especially for Figure 5e. This will help readers distinguish whether the green color in the image corresponds to cells or dye deposits.

We thank the reviewer for the valuable suggestion. We understand that nuclear staining can help determine the origin of fluorescence signals. However, in our three-dimensional imaging system, the deep signal acquisition range after tissue clearing often causes nuclear dyes such as DAPI to generate highly dense and widespread fluorescence, especially in regions rich in vascular structures, which can obscure the fine vascular and perivascular details of interest. Therefore, this study primarily focuses on high-resolution visualization of the spatial architecture of the vascular and biliary systems. We have added an explanation regarding this point in Figures S2I–J.

Reviewer #2 (Public review):

Summary:

The present manuscript of Xu et al. reports a novel clearing and imaging method focusing on the liver. The authors simultaneously visualized the portal vein, hepatic artery, central vein, and bile duct systems by injecting metal compound nanoparticles (MCNPs) with different colors into the portal vein, heart left ventricle, inferior vena cava, and the extrahepatic bile duct, respectively. The method involves: trans-cardiac perfusion with 4% PFA, the injection of MCNPs with different colors, clearing with the modified CUBIC method, cutting 200 micrometer thick slices by vibratome, and then microscopic imaging. The authors also perform various immunostaining (DAB or TSA signal amplification methods) on the tissue slices from MCNP-perfused tissue blocks. With the application of this methodical approach, the authors report dense and very fine vascular branches along the portal vein. The authors name them as 'periportal lamellar complex (PLC)' and report that PLC fine branches are directly connected to the sinusoids. The authors also claim that these structures co-localize with terminal bile duct branches and sympathetic nerve fibers, and contain endothelial cells with a distinct gene expression profile. Finally, the authors claim that PLC-s proliferate in liver fibrosis (CCl4 model) and act as a scaffold for proliferating bile ducts in ductular reaction and for ectopic parenchymal sympathetic nerve sprouting.

Strengths:

The simultaneous visualization of different hepatic vascular compartments and their combination with immunostaining is a potentially interesting novel methodological approach.

Weaknesses:

This reviewer has several concerns about the validity of the microscopic/morphological findings as well as the transcriptomics results. In this reviewer's opinion, the introduction contains overstatements regarding the potential of the method, there are severe caveats in the method descriptions, and several parts of the Results are not fully supported by the documentation. Thus, the conclusions of the paper may be critically viewed in their present form and may need reconsideration by the authors.

We sincerely thank the reviewer for the thorough evaluation and constructive comments on our study. We fully understand and appreciate the reviewer’s concerns regarding the methodological validity and interpretation of the results. In response, we have made comprehensive revisions and additions to the manuscript as follows:

First, we have carefully revised the Introduction and Discussion sections to provide a more balanced description of the methodological potential, removing statements that might be considered overstated, and clarifying the applicable scope and limitations of our approach (see the revised Introduction and Discussion).

Second, we have substantially expanded the Methods section with detailed information on model construction, imaging parameters, data processing workflow, and technical aspects of the single-cell transcriptomic reanalysis, to enhance the transparency and reproducibility of the study.

Third, we have added additional references and explanatory notes in the Results section to better support the main conclusions (see Section 6 of the Results).

Finally, we have rechecked and validated all experimental data, and conducted a verification analysis using an independent single-cell RNA-seq dataset (Figure S6). The results confirm that the morphological observations and transcriptomic findings are consistent and reproducible across independent experiments.

We believe these revisions have greatly strengthened the reliability of our conclusions and the overall scientific rigor of the manuscript. Once again, we sincerely appreciate the reviewer’s valuable comments, which have been very helpful in improving the logic and clarity of our work.

Reviewer #3 (Public review):

Summary:

In the reviewed manuscript, researchers aimed to overcome the obstacles of high-resolution imaging of intact liver tissue. They report successful modification of the existing CUBIC protocol into Liver-CUBIC, a high-resolution multiplex 3D imaging method that integrates multicolor metallic compound nanoparticle (MCNP) perfusion with optimized liver tissue clearing, significantly reducing clearing time and enabling simultaneous 3D visualization of the portal vein, hepatic artery, bile ducts, and central vein spatial networks in the mouse liver. Using this novel platform, the researchers describe a previously unrecognized perivascular structure they termed Periportal Lamellar Complex (PLC), regularly distributed along the portal vein axis. The PLC originates from the portal vein and is characterized by a unique population of CD34⁺Sca-1⁺ dual-positive endothelial cells. Using available scRNAseq data, the authors assessed the CD34⁺Sca-1⁺ cells' expression profile, highlighting the mRNA presence of genes linked to neurodevelopment, biliary function, and hematopoietic niche potential. Different aspects of this analysis were then addressed by protein staining of selected marker proteins in the mouse liver tissue. Next, the authors addressed how the PLC and biliary system react to CCL4-induced liver fibrosis, implying PLC dynamically extends, acting as a scaffold that guides the migration and expansion of terminal bile ducts and sympathetic nerve fibers into the hepatic parenchyma upon injury.

The work clearly demonstrates the usefulness of the Liver-CUBIC technique and the improvement of both resolution and complexity of the information, gained by simultaneous visualization of multiple vascular and biliary systems of the liver at the same time. The identification of PLC and the interpretation of its function represent an intriguing set of observations that will surely attract the attention of liver biologists as well as hepatologists; however, some claims need more thorough assessment by functional experimental approaches to decipher the functional molecules and the sequence of events before establishing the PLC as the key hub governing the activity of biliary, arterial, and neuronal liver systems. Similarly, the level of detail of the methods section does not appear to be sufficient to exactly recapitulate the performed experiments, which is of concern, given that the new technique is a cornerstone of the manuscript.

Nevertheless, the work does bring a clear new insight into the liver structure and functional units and greatly improves the methodological toolbox to study it even further, and thus fully deserves the attention of readers.

Strengths:

The authors clearly demonstrate an improved technique tailored to the visualization of the liver vasulo-biliary architecture in unprecedented resolution.

This work proposes a new biological framework between the portal vein, hepatic arteries, biliary tree, and intrahepatic innervation, centered at previously underappreciated protrusions of the portal veins - the Periportal Lamellar Complexes (PLCs).

Weaknesses:

Possible overinterpretation of the CD34+Sca1+ findings was built on re-analysis of one scRNAseq dataset.

Lack of detail in the materials and methods section greatly limits the usefulness of the new technique to other researchers.

We thank the reviewer for this important comment. We agree that when conclusions are mainly based on a single dataset, overinterpretation should be avoided. In response to this concern, we have carefully re-evaluated and clearly limited the scope of our interpretation of the scRNA-seq analysis. In addition, we performed a validation analysis using an independent single-cell RNA-seq dataset (see new Figure S6), which consistently confirmed the presence and characteristic transcriptional profile of the periportal CD34⁺Sca1⁺ endothelial cell population. These supplementary analyses strengthen the robustness of our findings and address the reviewer’s concern regarding potential overinterpretation.

In the revised manuscript, we have also greatly expanded the Materials and Methods section by providing detailed information on sample preparation, imaging parameters, data processing workflow, and single-cell reanalysis procedures. These revisions substantially improve the transparency and reproducibility of our methodology, thereby enhancing the usability and reference value of this technique for other researchers.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

Introduction

(1) In general, the Introduction is very lengthy and repetitive. It needs extensive shortening to a maximum of 2 A4 pages.

We thank the reviewer for the valuable suggestions. We have thoroughly condensed and restructured the Introduction, removing redundant content and merging related paragraphs to make the theme more focused and the logic clearer. The revised Introduction has been shortened to within two A4 pages, emphasizing the scientific question, innovation, and technical approach of the study.

(2) Please correct this erroneous sentence:

'...the liver has evolved the most complex and densely n organized vascular network in the body, consisting primarily of the portal vein system, central vein system, hepatic artery system, biliary system, and intrahepatic autonomic nerve network [6, 7].'

We thank the reviewer for pointing out this spelling error. The revised sentence is as follows:

“…the liver has evolved the most complex and densely organized ductal-vascular network in the body, consisting primarily of the portal vein system, central vein system, hepatic artery system, biliary system, and intrahepatic autonomic nerve network [6, 7].”

(3) '...we achieved a 63.89% improvement in clearing efficiency and a 20.12% increase in tissue transparency'

Please clarify what you exactly mean by 'clearing efficiency' and 'increased tissue transparency'.

We thank the reviewer for the valuable comments and have clarified the relevant terminology in the revised manuscript.

“Clearing efficiency” refers to the improvement in the time required for the liver tissue to become completely transparent when treated with the optimized Liver-CUBIC protocol (40% urea + H₂O₂), compared with the conventional CUBIC method. In this study, the clearing time was reduced from 9 days to 3.25 days, representing a 63.89% increase in time efficiency.

“Tissue transparency” refers to the ability of the cleared tissue to transmit visible light. We quantified the optical transparency by measuring light transmittance across the 400–900 nm wavelength range using a microplate reader. The results showed that the average transmittance increased by 20.12%, indicating that Liver-CUBIC treatment markedly enhanced the optical clarity of the liver tissue.

(4) I am concerned about claiming this imaging method as real '3D imaging'. Namely, while the authors clear full lobes, they actually cut the cleared lobes into 200-micrometer-thick slices and perform further microscopy imaging on these slices. Considering that they focus on ductular structures of the liver (such as vasculature, bile duct system, and innervations), 200 micrometer allows a very limited 3D overview, particularly in comparison with the whole-mount immuno-imaging methods combined with light sheet microscopy (such as Adori 2021, Liu 2021, etc). In this context, I feel several parts of the Introduction to be an overstatement: besides of emphasizing the advantages of the technique (such as simultaneous visualization of different hepatic vascular compartments and the bile duct system by MCNPs, the combination with immunostainings), the authors must honestly discuss the limitations (such as limited tissue overview, potential dye perfusion problems - uneven distribution of the dye etc).

We appreciate the reviewer’s insightful comments. It is true that most of the imaging depth in this study was limited to approximately 200 μm, and thus it could not achieve whole-liver three-dimensional imaging comparable to light-sheet microscopy. However, the primary focus of our study was to resolve the microscopic intrahepatic architecture, particularly the spatial relationships among blood vessels, bile ducts, and nerve fibers. Through high-resolution imaging of thick tissue sections, combined with MCNP-based multichannel labeling and immunofluorescence co-staining, we were able to accurately delineate the three-dimensional distribution of these microstructures within localized regions.

In addition to thick-section imaging, we also obtained whole-lobe dye perfusion data (as shown in Figure S1F), which comprehensively depict the three-dimensional branching patterns and distribution of the vascular systems within the liver lobe. These images were acquired from intact liver lobes perfused with MCNP dyes, revealing a continuous vascular network extending from major trunks to peripheral branches, thereby demonstrating that our approach is also capable of achieving organ-level visualization.

We have added this image and a corresponding description in the revised manuscript to more comprehensively present the coverage of our imaging system, and we have incorporated this clarification into the Discussion section.

Method

(5) More information may be needed about MCNPs:

a) As reported, there are nanoparticles with different colors in brightfield microscopy, but the particles are also excitable in fluorescence microscopy. Would you please provide a summary about excitation/emission wavelengths of the different MCNPs? This is crucial to understand to what extent the method is compatible with fluorescence immunohistochemistry.

We thank the reviewer for the careful attention and professional suggestion. We fully agree that this issue is critical for evaluating the compatibility of our method with fluorescent immunohistochemistry. Different types of metal compound nanoparticles (MCNPs) have clearly distinguishable spectral properties:

- MCNP-Green and MCNP-Yellow: AF488-matched spectra, with excitation/emission wavelengths of 495/519 nm.

- MCNP-Pink: Designed for far-red spectra, with excitation/emission wavelengths of 561/640 nm.

- MCNP-Black: Non-fluorescent, appearing black under bright-field microscopy only.

The above information has been added to the Materials and Methods section.

b) Also, is there more systematic information available concerning the advantage of these particles compared to 'traditional' fluorescence dyes, such as Alexa fluor or Cy-dyes, in fluorescence microscopy and concerning their compatibility with various tissue clearing methods (e.g., with the frequently used organic-solvent-based methods)?

We thank the reviewer for the detailed question. Compared with conventional organic fluorescent dyes, MCNP offers the following advantages:

- Enhanced photostability: Its inorganic core-shell structure resists fading even after hydrogen peroxide bleaching.

- High signal stability: Fluorescence is maintained during aqueous-based clearing (e.g., CUBIC) and multiple rounds of staining without quenching.

We appreciate the reviewer’s suggestion. In our Liver-CUBIC system, MCNP nanoparticles exhibited excellent multi-channel labeling stability and fluorescence signal retention. Regarding compatibility with other clearing methods (e.g., SCAFE, SeeDB, CUBIC), since these methods have limited effectiveness for whole-liver clearing (see Figure 2 of Tainaka, et al. 2014) and cannot meet the requirements for high-resolution microstructural imaging in this study, we consider further testing of their compatibility unnecessary.

In summary, MCNP dye demonstrates superior signal stability and spectral separation compared with conventional organic fluorescent dyes in multi-channel, long-term, high-transparency three-dimensional tissue imaging.

c) When you perfuse these particles, to which structures do they bind inside the ducts (vessels, bile ducts)? Is the 48h post-fixation enough to keep them inside the tubes/bind them to the vessel walls? Is there any 'wash-out' during the complex cutting/staining procedure? E.g., in Figure 2D: the 'classical' hepatic artery in the portal triad is not visible - but the MCNP apparently penetrated to the adjacent sinusoids at the edge of the lobulus. Also, in Figure 3B, there is a significant mismatch between the MNCP-green (bile duct) signal and the CD19 (epithelium marker) immunostaining. Please discuss these.

The experimental results showed that following injection, MCNP nanoparticles primarily remained within the vascular and biliary lumens, and their tissue distribution depended on physical perfusion. No dye signal was observed to diffuse into the surrounding parenchyma, nor did the particles adhere to cell surfaces or enter cells. The newly added Supplementary Figures S2A–H further confirm this feature: the dye signal is strictly confined within the lumens, clearly delineating continuous vascular paths and biliary branching patterns, strongly supporting the conclusion that “MCNP dye is distributed only within luminal spaces.”

Thus, MCNP dye mainly serves as an intraluminal tracer rather than a label for specific cell types.

We provide the following explanations and analyses regarding MCNP distribution in the hepatic vascular and biliary systems and its post-fixation stability:

- Potential signal displacement during sectioning/immunostaining: During slicing and immunostaining, a small number of particles may be washed away due to mechanical cutting or washing steps; however, the overall three-dimensional structure retains high spatial fidelity.

- Observation in Figure 2D: MCNP was seen entering the sinusoidal spaces at the lobule periphery, but hepatic arteries were not visible, likely due to limitations in section thickness. Although arteries were not apparent in this slice, arterial distribution around the portal vein is visible in Figure 2C. It should be noted that Figures 2C, D, and E do not represent whole-liver imaging, so not all regions necessarily contain visible hepatic arteries. For easier identification, the main hepatic artery trunk is highlighted in cyan in Figure 2E.

- Incomplete biliary signal in Figure 3B: This may be because CK19 labeling only covers biliary epithelial cells, whereas MCNP-green distributes throughout the biliary lumen. In Figure 3B, the terminal MCNP-green signal exhibits irregular polygonal structures, which we interpret as the canalicular regions.

(6) Which fixative was used for 48h of postfixation (step 6) after MCNP injections?

After MCNP injection, mouse livers were post-fixed in 4% paraformaldehyde (PFA) for 48 hours. This fixation condition effectively “locks” the MCNP particles within the vascular and biliary lumens, maintaining their spatial positions, while also being compatible with subsequent sectioning and multi-channel immunostaining analyses.

The above information has been added to the Materials and Methods section

(7) What is the 'desired thickness' in step 7? In the case of immunostained tissue, a 200-micrometer slice thickness is mentioned. However, based on the Methods, it is not completely clear what the actual thickness of the tissue was that was examined ultimately in the microscopes, and whether or not the clearing preceded the cutting or vice versa.

We appreciate the reviewer’s question. The “desired thickness” referred to in step 7 of the manuscript corresponds to the thickness of tissue sections used for immunostaining and high-resolution microscopic imaging, which is typically around 200 µm. We selected 200 µm because this thickness is sufficient to observe the PLC structure in its entirety, allows efficient staining, and preserves tissue architecture well. Other researchers may choose different section thicknesses according to their experimental needs.

In this study, the processing order for immunostained tissue samples was sectioning followed by clearing, as detailed below:

Section Thickness

To ensure antibody penetration and preservation of three-dimensional structure, tissue sections were typically cut to ~200 µm. Thicker sections can be used if more complete three-dimensional structures are required, but adjustments may be needed based on antibody penetration and fluorescence detection conditions.

Clearing Sequence

After sectioning, slices were processed using the Liver-CUBIC aqueous-based clearing system.

(8) More information is needed concerning the 'deep-focus microscopy' (Keyence), the applied confocal system, and the THUNDER 'high resolution imaging system': basic technical information, resolutions, objectives (N.A., working distance), lasers/illumination, filters, etc.

In this study, all liver lobes (left, right, caudate, and quadrate lobes) were subjected to Liver-CUBIC aqueous-based clearing to ensure uniform visualization of MCNP fluorescence and immunolabeling throughout the three-dimensional imaging of the entire liver.

The above information has been added to the Materials and Methods section.

Imaging Systems and Settings

VHX-6000 Extended Depth-of-Field Microscope: Objective: VH-Z100R, 100×–1000×; resolution: 1 µm (typical); illumination: coaxial reflected; transmitted illumination on platform: ON.

Zeiss Confocal Microscope (980): Objectives: 20× or 40×; image size: 1024 × 1024. Fluorescence detection was set up in three channels:

- Channel 1: 639 nm laser, excitation 650 nm, emission 673 nm, detection range 673–758 nm, corresponding to Cy5-T1 (red).

- Channel 2: 561 nm laser, excitation 548 nm, emission 561 nm, detection range 547–637 nm, corresponding to Cy3-T2 (orange).

- Channel 3: 488 nm laser, excitation 493 nm, emission 517 nm, detection range 490–529 nm, corresponding to AF488-T3 (green).

Leica THUNDER Imager 3D Tissue: Fluorescence detection in two channels:

- Channel 1: FITC channel (excitation 488 nm, emission ~520 nm).

- Channel 2: Orange-red channel (excitation/emission 561/640 nm).

Equipped with matching filter sets to ensure signal separation.

The above information has been added to the Materials and Methods section.

(9) Liver-CUBIC, step 2: which lobe(s) did you clear (...whole liver lobes...).

In this study, all liver lobes (left, right, caudate, and quadrate lobes) were subjected to Liver-CUBIC aqueous-based clearing to ensure uniform visualization of MCNP fluorescence and immunolabeling throughout the three-dimensional imaging of the entire liver.

The above information has been added to the Materials and Methods section.

(10) For the DAB and TSA IHC stainings, did you use free-floating slices, or did you mount the vibratome sections and do the staining on mounted sections?

In this study, fixed livers were first sectioned into thick slices (~200 µm) using a vibratome. Subsequently, DAB and TSA immunohistochemical (IHC) staining were performed on free-floating sections. During the entire staining process, the slices were kept floating in the solutions, ensuring thorough antibody penetration in the thick sections while preserving the three-dimensional tissue architecture, thereby facilitating multiple rounds of staining and three-dimensional imaging.

(11) Regarding the 'transmission quantification': this was measured on 1 mm thick slices. While it is interesting to make a comparison between different clearing methods in general, one must note that it is relatively easy to clear 1mm thick tissue slices with almost any kind of clearing technique and in any tissues. The 'real' differences come with thicker blocks, such as >5mm in the thinnest dimension. Do you have such experiences (e.g., comparison in whole 'left lateral liver lobes')?

In this study, we performed three-dimensional visualization of entire liver lobes to depict the distribution of MCNPs and the overall spatial architecture of the vascular and biliary systems (Figure S1F). However, due to the limitations of the plate reader and fluorescence imaging systems in terms of spatial resolution and light penetration depth, quantitative analyses were conducted only on tissue sections approximately 1 mm thick.

Regarding the comparative quantification of different clearing methods, as the reviewer noted, nearly all aqueous- or organic solvent–based clearing techniques can achieve relatively uniform transparency in 1 mm-thick tissue sections, so differences at this thickness are limited. We have not yet conducted systematic comparisons on whole-lobe sections thicker than 5 mm and therefore cannot provide “true” difference data for thicker tissues.

(12) There is no method description for the ELMI studies in the Methods.

Transmission Electron Microscopy (TEM) Analysis of MCNPs

Before imaging, the MCNP dye solution was centrifuged at 14,000 × g for 10 minutes at 4 °C to remove aggregates and impurities. The supernatant was collected, diluted 50-fold, and 3–4 μL of the sample was applied onto freshly glow-discharged Quantifoil R1.2/1.3 copper grids (Electron Microscopy Sciences, 300 mesh). The sample was allowed to sit for 30 seconds to enable particle adsorption, after which excess liquid was gently wicked away with filter paper and the grid was air-dried at room temperature. The sample was then negatively stained with 1% uranyl acetate for 30 seconds and air-dried again before imaging.

Negative-stain TEM images were acquired using a JEOL JEM-1400 transmission electron microscope operating at 120 kV and equipped with a CCD camera. Data acquisition followed standard imaging conditions.

The above information has been added to the Materials and Methods section.

(13) Please, provide a method description for the applied CCl4 cirrhosis model. This is completely missing.

(1) Under a fume hood, carbon tetrachloride (CCl₄) was dissolved in corn oil at a 1:3 volume ratio to prepare a working solution, which was filtered through a 0.2 μm filter into a 30 mL glass vial. In our laboratory, to mimic chronic injury, mice in the experimental group were intraperitoneally injected at a dose of 1 mL/kg body weight per administration.

(2) Mice were carefully removed from the cage and placed on a scale to record body weight for calculation of the injection volume.

(3) The needle cap was carefully removed, and the required volume of the pre-prepared CCl₄ solution was drawn into the syringe. The syringe was gently flicked to remove any air bubbles.

(4) Mice were placed on a textured surface (e.g., wire cage) and restrained. When the mouse was properly positioned, ideally with the head lowered about 30°, the left lower or right lower abdominal quadrant was identified.

(5) Holding the syringe at a 45° angle, with the bevel facing up, the needle was inserted approximately 4–5 mm into the abdominal wall, and the calculated volume of CCl₄ was injected.

(6) Mice were returned to their cage and observed for any signs of discomfort.

(7) Needles and syringes were disposed of in a sharps container without recapping. A new syringe or needle was used for each mouse.

(8) To establish a progressive liver fibrosis model, injections were administered twice per week (e.g., Monday and Thursday) for 3 or 6 consecutive weeks (n=3 per group). Control mice were injected with an equal volume of corn oil for 3 or 6 weeks (n=3 per group).

(9) Forty-eight hours after the last injection, mice were euthanized by cervical dislocation, and livers were rapidly harvested. Portions of the liver were processed for paraffin embedding and histological sectioning, while the remaining tissue was either immediately frozen or used for subsequent molecular biology analyses.

The above information has been added to the Materials and Methods section.

(14) Please provide a method description for the quantifications reported in Figures 5D, 5F, and 6E.

ImageJ software was used to analyze 3D stained images (Figs. 5F, 6E), and the ultra-depth-of-field 3D analysis module was used to analyze 3D DAB images (Fig. 5D). The specific steps are as follows:

Figure 5D: DAB-stained 3D images from the control group and the CCl4 6-week (CCl4-6W) group were analyzed. For each group, 20 terminal bile duct branch nodes were randomly selected, and the actual path distance along the branch to the nearest portal vein surface was measured. All measurements were plotted as scatter plots to reflect the spatial extension of bile ducts relative to the portal vein under different conditions.

Figure 5F: TSA 3D multiplex-stained images from the control group, CCl4 3-week (CCl4-3W), and CCl4 6-week (CCl4-6W) groups were analyzed. For each group, 5 terminal bile duct branch nodes were randomly selected, and the actual path distance along the branch to the nearest portal vein surface was measured. Measurements were plotted as scatter plots to illustrate bile duct spatial extension.

Figure 6E: TSA 3D multiplex-stained images from the control, CCl4-3W, and CCl4-6W groups were analyzed. For each group, 5 terminal nerve branch nodes were randomly selected, and the actual path distance along the branch to the nearest portal vein surface was measured. Scatter plots were generated to depict the spatial distribution of nerves under different treatment conditions.

(15) Please provide a method description for the human liver samples you used in Figure S6. Patient data, fixation, etc...

The human liver tissue samples shown in Figure S6 were obtained from adjacent non-tumor liver tissues resected during surgical operations at West China Hospital, Sichuan University. All samples used were anonymized archived tissues, which were applied for scientific research in accordance with institutional ethical guidelines and did not involve any identifiable patient information. After being fixed in 10% neutral formalin for 24 hours, the tissues were routinely processed for paraffin embedding (FFPE), and sectioned into 4 μm-thick slices for immunostaining and fluorescence imaging.

Results

(16) While it is stated in the Methods that certain color MCNPs were used for labelling different structures (i.e., yellow: hepatic artery; green: bile duct; portal vein: pink; central veins: black), in some figures, apparently different color MCNPs are used for the respective structures. E.g., in Figure 1J, the artery is pink and the portal vein is green. Please clarify this.

The color assignment of MCNP dyes is not fixed across different experiments or schematic illustrations. MCNP dyes of different colors are fundamentally identical in their physical and chemical properties and do not exhibit specific binding or affinity for particular vascular structures. We select different colors based on experimental design and imaging presentation needs to facilitate distinction and visualization, thereby enhancing recognition in 3D reconstruction and image display. Therefore, the color labeling in Figure 1F is primarily intended to illustrate the distribution of different vascular systems, rather than indicating a fixed correspondence to a specific dye or injection color.

(17) In Figure 1J, the hepatic artery is extremely shrunk, while the portal vein is extremely dilated - compared to the physiological situation. Does it relate to the perfusion conditions?

We appreciate the reviewer’s attention. In fact, under normal physiological conditions, the hepatic arteries labeled by CD31 are naturally narrow. Therefore, the relatively thin hepatic arteries and thicker portal veins shown in Figure 1J are normal and unrelated to the perfusion conditions. See figure 1E of Adori et al., 2021.

(18) Re: MCNP-black labelled 'oval fenestrae': the Results state 50-100 nm, while they are apparently 5-10-micron diameter in Figure 1I. Accordingly, the comparison with the ELMI studies in the subsequent paragraph is inappropriate.

We thank the reviewer for the correction. The previous statement was a typographical error. In fact, the diameter of the “elliptical windows” marked by MCNP-black is 5–10 μm, so the diameter of 5–10 μm shown in Figure 1I is correct.

(19) Please, correct this erroneous sentence: 'Pink marked the hepatic arterial system by injection extrahepatic duct (Figure 2B).'

Original sentence: “The hepatic arterial system was labeled in pink by injection through the extrahepatic duct (Figure 2B).”

Revised sentence: “The hepatic arterial system was labeled in pink by injection through the left ventricle (Figure 2B).”

(20) How do you define the 'primary portal vein tract'?

We thank the reviewer for the question. The term “primary portal vein tract” refers to the first-order branches of the portal vein that enter the liver from the hepatic hilum. These are the major branches arising directly from the main portal vein trunk and are responsible for supplying blood to the respective hepatic lobes. This definition corresponds to the concept of the first-order portal vein in hepatic anatomy.

(21) I am concerned that the 'periportal lamellar complex (PLC)' that the Authors describe really exists as a distinct anatomical or functional unit. I also see these in 3D scans - in my opinion, these are fine, lower-order portal vein branches that connect the portal veins to the adjacent sinusoid. The strong MCNP-labelling of these structures may be caused by the 'sticking' of the perfused MCNP solutions in these 'pockets' during the perfusion process. What do these structures look like with SMA or CD31 immunostaining? Also, one may consider that the anatomical evaluation of these structures may have limitations in tissue slices. Have you ever checked MCNP-perfused, cleared full live lobes in light sheet microscope scans? I think this would be very useful to have a comprehensive morphological overview. Unfortunately, based on the presented documentation, I am also not convinced that PLCs are 'co-localize' with fine terminal bile duct branches (Figure 3E, S3C), or with TH+ 'neuronal bead chain networks' (Fig 6C). More detailed and more convincing documentation is needed here.

We thank the reviewer for the detailed comments. Regarding the existence and function of the periportal lamellar complex (PLC), our observations are based on MCNP-Pink labeling of the portal vein, through which we were able to identify the PLC structure surrounding the portal branches. It should be noted that the PLC represents a very small anatomical structure. Although we have not yet performed light-sheet microscopy scanning, we anticipate that such imaging would primarily visualize larger portal vein branches. Nevertheless, this does not affect our overall conclusions.

We also appreciate the reviewer’s suggestion that the observed structures might result from MCNP adherence during perfusion. To verify the structural characteristics of the PLC, we performed immunostaining for SMA and CD31, which revealed a specific arrangement pattern of smooth muscle and endothelial markers rather than simple perfusion-induced deposition (Figures 4F and S6B).

Regarding the apparent colocalization of the PLC with terminal bile duct branches (Figures 3E and S3C) and TH⁺ neuronal bead-like networks (Figure 6C), we acknowledge that current literature evidence remains limited. Therefore, we have carefully described these observations as possible spatial associations rather than definitive conclusions. Future studies integrating high-resolution three-dimensional imaging with functional analyses will help to further clarify the anatomical and physiological significance of the PLC.

(22) 'Extended depth-of-field three-dimensional bright-field imaging revealed a strict 1:1 anatomical association between the primary portal vein trunk (diameter 280 {plus minus} 32 μm) and the first-order bile duct (diameter 69 {plus minus} 8 μm) (Figures 3A and S3A)'.

How do you define '1:1 anatomical association'? How do you define and identify the 'order' (primary, secondary) of vessel and bile duct branches in 200-micrometer slices?

We thank the reviewer for the question. In this study, the term “1:1 anatomical correlation” refers to the stable paired spatial relationship between the main portal vein trunk and its corresponding primary bile duct within the same portal territory. In other words, each main portal vein branch is accompanied by a primary bile duct of matching branching order and trajectory, together forming a “vascular–biliary bundle.”

The definitions of “primary” and “secondary” branches were based on extended-depth 3D bright-field reconstructions, considering both branching hierarchy and vessel/duct diameters: primary branches arise directly from the main trunk at the hepatic hilum and exhibit the largest diameters (averaging 280 ± 32 μm for the portal vein and 69 ± 8 μm for the bile duct), whereas secondary branches extend from the primary branches toward the lobular interior with smaller calibers.

(23) In my opinion, the applied methodical approach in the single cell transcriptomics part (data mining in the existing liver single cell database and performing Venn diagram intersection analysis in hepatic endothelial subpopulations) is largely inappropriate and thus, all the statements here are purely speculative. In my opinion, to identify the molecular characteristics of such small and spatially highly organized structures like those fine radial portal branches, the only way is to perform high-resolution spatial transcriptomic.

We thank the reviewer for the comment. We fully acknowledge the importance of high-resolution spatial transcriptomics in identifying the fine structural characteristics of portal vein branches. Due to current funding and technical limitations, we were unable to perform such high-resolution spatial transcriptomic analyses. However, we validated the molecular features of the PLC using another publicly available liver single-cell RNA-sequencing dataset, which provided preliminary supporting evidence (Figures S6B and S6C). In the manuscript, we have carefully stated that this analysis is exploratory in nature and have avoided overinterpretation. In future studies, high-resolution spatial omics approaches will be invaluable for more precisely delineating the molecular characteristics of these fine structures.

(24) 'How the autonomic nervous system regulates liver function in mice despite the apparent absence of substantive nerve fiber invasion into the parenchyma remains unclear.'

Please consider the role of gap junctions between hepatocytes (e.g., Miyashita, 1991; Seseke, 1992).

In this study, we analyzed the spatial distribution of hepatic nerves in mice using immunofluorescence staining and found that nerve fibers were almost exclusively confined to the portal vein region (Figure S6A). Notably, this distribution pattern differs markedly from that in humans. Previous studies have shown that, in human livers, nerves are not only located around the portal veins but also present along the central veins, interlobular septa, and within the parenchymal connective tissue (Miller et al., 2021; Yi, la Fleur, Fliers & Kalsbeek, 2010).

Further research has provided a physiological explanation for this interspecies difference: even among species with distinct sympathetic innervation patterns in the parenchyma—i.e., with or without direct sympathetic input—the sympathetic efferent regulatory functions may remain comparable (Beckh, Fuchs, Ballé & Jungermann, 1990). This is because signals released from aminergic and peptidergic nerve terminals can be transmitted to hepatocytes through gap junctions as electrical signals (Hertzberg & Gilula, 1979; Jensen, Alpini & Glaser, 2013; Seseke, Gardemann & Jungermann, 1992; Taher, Farr & Adeli, 2017).

However, the scarcity of nerve fibers within the mouse hepatic parenchyma suggests that the mechanisms by which the autonomic nervous system regulates liver function in mice may differ from those in humans. This observation prompted us to further investigate the potential role of PLC endothelial cells in this process.

(25) Please, correct typos throughout the text.

We thank the reviewer for this comment. We have carefully proofread the entire manuscript and corrected all typographical errors and minor language issues throughout the text.

Reviewer #3 (Recommendations for the authors):

(1) A strong recommendation - the authors ought to challenge their scRNAsq- re-analysis with another scRNAseq dataset, namely a recently published atlas of adult liver endothelial, but also mesenchymal, immune, and parenchymal cell populations https://pubmed.ncbi.nlm.nih.gov/40954217/, performed with Smart-seq2 approach, which is perfectly suitable as it brings higher resolution data, and extensive cluster identity validation with stainings. Pietilä et al. indicate a clear distinction of portal vein endothelial cells into two populations that express Adgrg6, Jag1 (e2c), from Vegfc double-positive populations (e5c and e2c). Moreover, the dataset also includes the arterial endothelial cells that were shown to be part of the PLC, but were not followed up with the scRNAseq analysis. This distinction could help the authors to further validate their results, better controlling for cross-contaminations that may occur during scRNAseq preparation.

We thank the reviewer for the valuable suggestion. As noted, we have further validated the molecular characteristics of the PLC using a recently published atlas of adult liver endothelial cells (Pietilä et al., 2023, PMID: 40954217). This dataset, generated using the Smart-seq2 technique, provides high-resolution transcriptomic profiles. By analyzing this dataset, we identified a CD34⁺LY6A⁺ portal vein endothelial cell population within the e2 cluster, which is localized around the portal vein. We then examined pathways and gene expression patterns related to hematopoiesis, bile duct formation, and neural signaling within these cells. The results revealed gene enrichment patterns consistent with those observed in our primary dataset, further supporting the robustness of our analysis of the PLC’s molecular characteristics.

(2) Improving the methods section is highly recommended, this includes more detailed information for material and protocols used - catalog numbers; protocol details of the usage - rocking platforms, timing, and tubes used for incubations; GitHub or similar page with code used for the scRNA seq re-analysis.

We thank the reviewer for the valuable suggestion. We have added more detailed information regarding the materials and experimental procedures in the Methods section, including catalog numbers, incubation conditions (such as the type of shaker, incubation time, and tube specifications), and other relevant parameters.

(3) In Figure 2A, the authors claim the size of the nanoparticle is 100nm, while based on the image, the size is ~150-180nm. A more thorough quantification of the particle size would help users estimate the usability of their method for further applications.

We thank the reviewer for the comment. In the TEM image shown in Figure 2A, the nanoparticles indeed appear to be approximately 150–200 nm in size. We have re-verified the particle dimensions and will update the corresponding description in the Methods section to allow readers to more accurately assess the applicability of this approach.

(4) In Figure 3E, it is not clear what is labeled by the pink signal. Please consider labeling the structures in the figure.

We thank the reviewer for the valuable comment. The pink signal in Figure 3E was originally intended to label the hepatic artery. However, a slight spatial misalignment occurred during the labeling process, making its position appear closer to the central vein rather than the portal vein in the image. To avoid misunderstanding, we will add clear annotations to the image and clarify this deviation in the figure legend in the revised version. It should also be noted that this figure primarily aims to illustrate the spatial relationship between the bile duct and the portal vein, and this minor deviation does not affect the reliability of our experimental conclusions.

(5) The following statement is not backed by quantification as it ought to be „Dual-channel three-dimensional confocal imaging combined with CK19 immunostaining revealed that the sites of dye leakage did not coincide with the CK19-positive terminal bile duct epithelium, but instead were predominantly localized within regions adjacent to the PLC structures".

We thank the reviewer for the valuable comment. We have added the corresponding quantitative analysis to support this conclusion. Quantitative assessment of the extended-depth imaging data revealed that dye leakage predominantly occurred in regions adjacent to the PLC structure, rather than in the perivenous sinusoidal areas. The corresponding results have been presented in the revised Figure 3G.

(6) Similarly, Figure 4F is central to the Sca1CD34 cell type identification but lacks any quantification, providing it would strengthen the key statement of the article. A possible way to approach this is also by FACS sorting the double-positive cells and bluk/qRT validation.

We thank the reviewer for raising this point. We agree that quantitative validation of the Sca1⁺CD34⁺ population by FACS sorting could further support our conclusions. However, the primary focus of this study is on the spatial localization and transcriptional features of PLC endothelial cells. The identification of the Sca1⁺CD34⁺ subset is robustly supported by multiple complementary approaches, including three-dimensional imaging, co-staining with pan-endothelial markers, and projection mapping analyses. Collectively, these lines of evidence provide a solid basis for characterizing this unique endothelial population.

(7) The images in Figure S4D are not comparable, as the Sca1-stained image shows a longitudinal section of the PV, but the other stainings are cross-sections of PVs.

We thank the reviewer for the careful comment. We agree that the original Sca1-stained image, being a longitudinal section of the portal vein, was not optimal for direct comparison with other cross-sectional images. We have replaced it with a cross-sectional image of the portal vein to ensure comparability across all images. The updated image has been included in the revised Supplementary Figure S4D.

(8) I might be wrong, but Figure 4J is entirely missing, and only a cartoon is provided. Either remove the results part or provide the data.

We appreciate the reviewer’s careful observation. Figure 4J was intentionally designed as a schematic illustration to summarize the structural relationships and spatial organization of the portal vein, hepatic artery, and PLC identified in the previous panels (Figures 4A–4I). It does not represent newly acquired experimental data, but rather serves to provide a conceptual overview of the findings.

To avoid misunderstanding, we have clarified this point in the figure legend and the main text, stating that Figure 4J is a schematic summary rather than an experimental image. Therefore, we respectfully prefer to retain the schematic figure to aid readers’ interpretation of the preceding results.

(9) The methods section lacks information about the CCL4concentration, and it is thus hard to estimate the dosage of CCL4 received (ml/kg). This is important for the interpretation of the severity of the fibrosis and presence of cirrhosis, as different doses may or may not lead to cirrhosis within the short regimen performed by the authors [PMID: 16015684 DOI: 10.3748/wjg.v11.i27.4167]. Validation of the fibrosis/cirrhosis severity is, in this case, crucial for the correct interpretation of the results. If the level of cirrhosis is not confirmed, only progressive fibrosis should be mentioned in the manuscript, as these two terms cannot be used interchangeably.

Thank you for the reviewer’s comment. We indeed omitted the information on the concentration of carbon tetrachloride (CCl4) in the Methods section. In our experiments, mice received intraperitoneal injections of CCl4 at a dose of 1 mL/kg body weight, twice per week, for a total of six weeks. We have revised the manuscript accordingly, using the term “progressive fibrosis” to avoid confusion between fibrosis and cirrhosis.

(10) The following statement is not backed by any correlation analysis: "Particularly during liver fibrosis progression, the PLC exhibits dynamic structural extension correlating with fibrosis severity,.. ".

We thank the reviewer for the comment. The original statement that the “PLC correlates with fibrosis severity” lacked support from quantitative analysis. To ensure a precise description, we have revised the sentence as follows: “During liver fibrosis progression, the PLC exhibits dynamic structural extension.”

(11) Similarly, the following statement is not followed by data that would address the impact of innervation on liver function: "How the autonomic nervous system regulates liver function in mice despite the apparent absence of substantive nerve fiber invasion into the parenchyma remains unclear.".

This section has been revised. In this study, we analyzed the spatial distribution of nerves in the mouse liver using immunofluorescence staining. The results showed that nerve fibers were almost entirely confined to the portal vein region (Figure S6A). Notably, this distribution pattern differs significantly from that in humans. Previous studies have demonstrated that in the human liver, nerves are not only distributed around the portal vein but also present in the central vein, interlobular septa, and connective tissue of the hepatic parenchyma (Miller et al., 2021; Yi, la Fleur, Fliers & Kalsbeek, 2010).

Previous studies have further explained the physiological basis for this difference: even among species with differences in parenchymal sympathetic innervation (i.e., species with or without direct sympathetic input), their sympathetic efferent regulatory functions may still be similar (Beckh, Fuchs, Ballé & Jungermann, 1990). This is because signals released by adrenergic and peptidergic nerve terminals can be transmitted to hepatocytes as electrical signals through intercellular gap junctions (Hertzberg & Gilula, 1979; Jensen, Alpini & Glaser, 2013; Seseke, Gardemann & Jungermann, 1992; Taher, Farr & Adeli, 2017). However, the scarcity of nerve fibers in the mouse hepatic parenchyma suggests that the mechanism by which the autonomic nervous system regulates liver function in mice may differ from that in humans. This finding also prompts us to further explore the potential role of PLC endothelial cells in this process.

(12) Could the authors discuss their interpretation of the results in light of the fact that the innervation is lower in cirrhotic patients? https://pmc.ncbi.nlm.nih.gov/articles/PMC2871629/. Also, while ADGRG6 (Gpr126) may play important roles in liver Schwann cells, it is likely not through affecting myelination of the nerves, as the liver nerves are not myelinated https://pubmed.ncbi.nlm.nih.gov/2407769/ and https://www.pnas.org/doi/10.1073/pnas.93.23.13280.

We have revised the text to state that although most hepatic nerves are unmyelinated, GPR126 (ADGRG6) may regulate hepatic nerve distribution via non-myelination-dependent mechanisms. Studies have shown that GPR126 exerts both Schwann cell–dependent and –independent functions during peripheral nerve repair, influencing axon guidance, mechanosensation, and ECM remodeling (Mogha et al., 2016; Monk et al., 2011; Paavola et al., 2014).

(13) The manuscript would benefit from text curation that would:

a) Unify the language describing the PLC, so it is clear that (if) it represents protrusions of the portal veins.

We have standardized the description of the PLC throughout the manuscript, clearly specifying its anatomical relationship with the portal vein. Wherever appropriate, we indicate that the PLC represents protrusions associated with the portal vein, avoiding ambiguous or inconsistent statements.

b) Increase the accuracy of the statements.

Examples: "bile ducts, and the central vein in adult mouse livers."

We have refined all statements for accuracy.

c) Reduce the space given to discussion and results in the introduction, moving them to the respective parts. The same applies to the results section, where discussion occurs at more places than in the Discussion part itself.

We have edited the Introduction, removing detailed results and functional explanations, and retaining only a concise overview.

Examples: "The formation of PLC structures in the adventitial layer may participate in local blood flow regulation, maintenance of microenvironmental homeostasis, and vascular-stem cell interactions."

"This finding suggests that PLC endothelial cells not only regulate the periportal microcirculatory blood flow, but also establish a specialized microenvironment that supports periportal hematopoietic regulation, contributing to stem cell recruitment, vascular homeostasis, and tissue repair. "

"Together, these findings suggest the PLC endothelium may act as a key regulator of bile duct branching and fibrotic microenvironment remodeling in liver cirrhosis. " This one in particular would require further validation with protein stainings and similar, directly in your model.

d) Provide a clear reference for the used scRNA seq so it's clear that the data were re-analyzed.

Example: "single-cell transcriptomic analysis revealed significant upregulation of bile duct-related genes in the CD34+Sca-1+ endothelium of PLC in cirrhotic liver, with notably high expression of Lgals1 (Galectin-1) and HGF(Figure 5G) "

When describing the transcriptional analysis of PLC endothelial cells, we explicitly cited the original scRNA-seq dataset (Su et al., 2021), clarifying that these data were reanalyzed rather than newly generated.

e) Introducing references for claims that, in places, are crucial for further interpretation of experiments.

Examples: "It not only guides bile duct branching during development but also"; the authors show no data from liver development.

Thank you for pointing this out. We have revised the relevant statement to ensure that the claim is accurate and well-supported.

f) Results sentence "Instead, bile duct epithelial cells at the terminal ducts extended partially along the canalicular network without directly participating in the formation of the bile duct lumen." Lacks a callout to the respective Figure.

We would like to thank the reviewers for pointing out this issue. In the revised manuscript, the relevant image (Figure 3D) has been clearly annotated with white arrows to indicate the phenomenon of terminal cholangiocytes extending along the bile canaliculi network. Additionally, the schematic diagram on the right side clearly shows the bile canaliculi, cholangiocytes, and bile flow direction using arrows and color coding, thus intuitively corresponding to the textual description.

(14) Formal text suggestions: The manuscript text contains a lot of missed or excessive spaces and several typos that ought to be fixed. A few examples follow:

a) "densely n organized vascular network "

b) "analysis, while offering high spatial "

c) "specific differences, In the human liver, "

d) Figure 4F has a typo in the description.

e) "generation of high signal-to-noise ratio, multi-target " SNR abbreviation was introduced earlier.

f) Canals of Hering, CoH abbreviation comes much later than the first mention of the Canals of Hering.

We thank the reviewer for the helpful comment regarding textual consistency. We have carefully reviewed and revised the entire manuscript to improve the accuracy, clarity, and consistency of the text.

Changes from the earlier version to the revised version

Abstract

Added in the revised version: The liver is a complex organ responsible for multiple functions, including metabolism, energy storage, detoxification, bile secretion, and immune regulation. Its highly organized vascular system plays a crucial role in maintaining functional zonation and tissue homeostasis. Within the liver, the hepatic artery, portal vein, hepatic vein, bile duct, and nerve networks intertwine to form an intricate three-dimensional architecture; however, traditional two-dimensional imaging fails to reveal their true spatial relationships, and current three-dimensional imaging methods remain insufficient to capture fine structural details. To achieve comprehensive visualization of these multi-ductal systems, we established a high-resolution three-dimensional imaging platform that combines multicolor perfusion of metallic compound nanoparticles (MCNPs) with an optimized tissue-clearing protocol (Liver-CUBIC), enabling simultaneous 3D reconstruction of the portal vein, hepatic artery, bile duct, and hepatic vein in mouse livers. Based on these data, we identified and defined a previously unrecognized structure located in the outer layer of the portal vein, termed the Periportal Lamellar Complex (PLC). The PLC encircles the portal vein between the vascular endothelium and the perisinusoidal region, exhibits low-permeability barrier characteristics, and contains a distinctive population of CD34+Sca-1+ endothelial cells. During liver fibrosis, the PLC extends from the portal vein toward the hepatic lobule, forming a structural scaffold that guides bile duct and nerve migration.

Removed in the revised version: The liver’s microenvironment consists of interconnected vascular, biliary, and neural networks that regulate homeostasis and disease progression. However, the lack of high-resolution 3D visualization hinders our understanding of their interactions in health and pathology.

Removed in the revised version: In this study, we developed a high-resolution multiplex 3D imaging method integrating multicolor metallic compound nanoparticle (MCNP) perfusion and optimized CUBIC tissue clearing, enabling simultaneous 3D visualization of the portal vein, hepatic artery, bile ducts, and central vein spatial networks in mouse liver.

Removed in the revised version: Using this approach, we identified a novel perivascular structure, the Periportal Lamellar Complex (PLC), regularly distributed along the portal vein axis between its endothelium and the periportal lobule sinusoid. The PLC contains a unique population of CD34⁺Sca-1⁺ dual-positive endothelial cells with a distinct gene expression profile, potentially linked to bile acid transport, biliary epithelial cell proliferation, and neural development.

Removed in the revised version: Under normal physiological conditions, the PLC colocalizes with terminal bile duct branches and sympathetic nerve fibers. During liver fibrosis progression, the PLC dynamically extends into liver lobules alongside advancing fibrosis, acting as a scaffold guiding the migration of bile ducts and nerve fibers into the hepatic parenchyma.

Removed in the revised version: Differential gene expression and GO enrichment analyses further revealed neurodevelopmental and axon guidance pathway enrichment in the CD34⁺Sca-1⁺ endothelial population of PLC during fibrosis.

Removed in the revised version: Collectively, these findings establish the PLC as a critical structural and functional hub integrating hepatic vascular, biliary, and neural systems to regulate liver function. This study advances our understanding of liver biology and offers potential therapeutic targets for liver diseases.

Introduction

Earlier version: The liver performs diverse physiological functions including nutrient metabolism, detoxification, bile secretion, and immune regulation, all of which rely on its unique microanatomy — the hepatic lobule, the smallest structural and functional unit [1–5]. To sustain the efficient operation of the vast number of hepatic lobules distributed throughout the parenchyma, the liver has evolved the most complex and densely n organized vascular network in the body, consisting primarily of the portal vein system, central vein system, hepatic artery system, biliary system, and intrahepatic autonomic nerve network [6, 7]. Together, these systems form a multicompartmental 3D network, integrating blood transport, bile drainage, and neuroimmune regulation essential for maintaining the hepatic microenvironment. Despite their distinct physiological functions, these systems are closely interconnected, coordinating with one another to support metabolic, immune, and regenerative processes within each lobule [7]. While macroscopic hepatic vascular patterns and metabolic zonation have been described [8, 9]. how these networks interact with countless microscale, spatially heterogeneous lobules remains poorly understood.

Revised version: The liver performs diverse physiological functions, including nutrient metabolism, detoxification, bile secretion, and immune regulation, all of which depend on its unique microanatomy—the hepatic lobule, the smallest structural and functional unit (Almazroo, Miah, & Venkataramanan, 2017; Friedman, 2008; Kalra et al., 2023; Ozougwu, 2017; Yap & Mitchell, 2022). To maintain efficient operation across the vast hepatic parenchyma, the liver has evolved an exceptionally complex ductal-vascular network composed of the portal vein, central vein, hepatic artery, biliary system, and intrahepatic autonomic nerve network (Miller, Oderberg, & Goessling, 2021; Q. Zhang et al., 2022). These systems form a multicompartmental 3D network that integrates blood transport, bile drainage, and neuroimmune regulation essential for maintaining hepatic homeostasis. Despite their distinct roles, these systems are spatially and functionally intertwined to coordinate metabolism, immunity, and regeneration within each lobule (Q. Zhang et al., 2022). While macroscopic vascular patterns and metabolic zonation have been described (Lv et al., 2021; Peeters et al., 2018), the micro-level interactions among these ductal systems within individual lobules remain poorly understood.

Added in the revised version: Traditional two-dimensional (2D) histology provides high spatial resolution but cannot reconstruct the true 3D continuity of vascular structures or their relationships across multiple ductal systems (Hankeova et al., 2021; Kim, Chen, Alpini, & Lau, 2015). Thus, three-dimensional (3D) visualization of multiple intrahepatic networks is essential for understanding liver development, homeostasis, and disease. Conventional imaging modalities such as ultrasound, micro-computed tomography (micro-CT), and magnetic resonance imaging (MRI) offer macroscopic resolution but cannot resolve micrometer-scale vascular details. X-ray phase-contrast computed tomography (PCCT) achieves contrast-agent-free microvascular imaging in rat livers (Qin et al., 2017), yet still lacks subcellular resolution. Similarly, the double resin-casting micro computed tomography (DUCT) technique reconstructed 3D portal vein–bile duct networks using dual-color resin casting, but failed to visualize luminal structures smaller than 5 μm (Hankeova et al., 2021). Therefore, a high-resolution, multicolor 3D imaging approach is required to elucidate terminal interactions such as portal vein–artery coupling, bile duct–vessel countercurrent exchange, and nerve–vessel communication.

Added in the revised version: Developing such imaging approaches for the liver remains technically challenging. The hepatic vasculature is highly dense (sinusoids occupy 10–15% of the liver volume) and intricately intertwined, demanding submicron resolution and millimeter-scale penetration depth (Lv et al., 2021; Segovia-Miranda et al., 2019). Moreover, tissue-specific autofluorescence (e.g., from bilirubin, lipofuscin, collagen, and NADPH/FAD) impairs optical clearing and fluorescence signal quality, yielding low signal-to-noise ratios (SNR) (Molina et al., 2022; Susaki et al., 2014; Tainaka et al., 2014). Finally, the liver’s multiple intrahepatic systems—blood vessels, bile ducts, nerves, lymphatics, and immune networks—are functionally interdependent, yet current methods lack efficient 3D multi-target labeling strategies to visualize them simultaneously. Consequently, existing clearing and labeling protocols (such as CLARITY, iDISCO, and conventional immunofluorescence) remain limited in depth, resolution, and SNR, leaving the liver far behind other organs in 3D multicolor imaging (Liu et al., 2022; Molina et al., 2022).

Added in the revised version: To address these challenges, we developed a liver-adapted, high-fidelity 3D visualization strategy tailored for the hepatic multi-ductal system. This platform integrates several key innovations: (1) The optimized Liver-CUBIC tissue clearing protocol shortened the clearing duration by 63.89% and enhanced optical transmittance by 20.12%, compared with the original CUBIC protocol; (2) Multicolor metal compound nanoparticles (MCNPs), which provide stable, simultaneous labeling of the portal vein, hepatic artery, bile duct, and central vein, and support multimodal imaging at subcellular resolution; and (3) A multimodal immunolabeling system compatible with 3,3’-diaminobenzidine (DAB) and tyramide signal amplification (TSA)-based multiplex staining, allowing high signal-to-noise 3D mapping of vascular and cellular structures.

Added in the revised version: Using this platform, we achieved liver 3D visualization at single-cell resolution and discovered a previously unrecognized structure, the Periportal Lamellar Complex (PLC). The PLC is periodically distributed along the portal vein axis and closely associated with terminal bile duct branches and autonomic nerve plexuses, forming a specialized interaction unit.

Added in the revised version: Single-cell transcriptomic and multiplex 3D immunostaining analyses revealed that the PLC exhibits distinct localization, morphology, and gene expression signatures, suggesting it may play a crucial role in hepatic microanatomical coordination.

Removed in the revised version: Traditional two-dimensional (2D) histological analysis, while offering high spatial resolution, has long been regarded as the gold standard for morphological studies [10, 11]. However, it presents notable limitations when applied to the investigation of complex three-dimensional (3D) vascular systems. On one hand, single-section analyses cannot restore the full 3D connectivity of vascular structures, while serial section reconstruction is prone to alignment errors and remains inadequate for resolving the interactions among multiple ductal systems. Therefore, achieving 3D-level visualization of multiple intrahepatic vascular systems is essential for studying their dynamic interactions during liver development, homeostasis, and disease progression.

Removed in the revised version: Conventional imaging modalities such as ultrasound, micro-computed tomography (Micro-CT), and magnetic resonance imaging (MRI) are widely applied in both experimental animal studies and clinical practice. Although these techniques can capture gross pathological alterations in the hepatic vasculature, they lack the spatial resolution required to resolve micrometer-scale structures. X-ray phase-contrast computed tomography (PCCT), recently developed for liver microvascular imaging, has achieved contrast-agent-free microvasculature visualization in rat livers[12] . However, its contrast sensitivity remains insufficient for resolving subcellular structures. Similarly, the DUCT technique successfully reconstructed a 3D portal vein–bile duct network in the mouse liver via dual-color resin casting, but was limited by its inability to resolve luminal structures smaller than 5 μm [11]. To fully understand how multiple hepatic ductal systems maintain their intricate terminal 3D interactions — such as portal vein– hepatic artery blood flow coupling, bile duct–vessel countercurrent exchange, and nerve–vessel signaling — and how these systems collectively coordinate with liver lobules, a more effective multicolor, high-resolution 3D imaging strategy is urgently needed. Nevertheless, integrating such approaches for the liver remains challenging for several reasons. First, the hepatic vasculature is structurally complex, featuring ultra- dense vascular distributions (with hepatic sinusoids occupying 10–15% of liver volume) and multilevel, intertwined arrangements (arterial branches encircling portal vein branches, and bile ducts running parallel to blood vessels). Imaging must therefore achieve submicron spatial resolution (≤1 μm) while maintaining millimeter-scale tissue penetration depth[8, 13].Second, the liver’s tissue-specific properties significantly interfere with imaging. Endogenous pigments (such as bilirubin and lipofuscin) and strong autofluorescent components (such as collagen and NADPH/FAD) severely impair tissue-clearing efficiency and degrade fluorescence signal quality, resulting in a low signal-to-noise ratio (SNR) in traditional clearing–staining–imaging workflows[14–16].Third, multiple intrahepatic systems — including blood vessels, bile ducts, nerves, lymphatics, and immune cell networks — are densely interwoven and functionally interdependent. However, there is currently a lack of efficient 3D multi-target labeling strategies capable of simultaneously visualizing these systems in both their hierarchical macroscopic structures and their fine terminal interactions. In summary, the liver’s unique structural complexity and tissue-specific properties impose significant constraints on the performance of existing clearing techniques (such as CLARITY and iDISCO) and labeling protocols (such as immunohistochemistry and immunofluorescence) in 3D imaging applications, limiting imaging depth, resolution, and SNR[14, 17]. Compared to organs such as the brain and kidney, the application of 3D, high-resolution, multicolor imaging in the liver remains markedly underdeveloped.

Removed in the revised version: Given these limitations, the development of liver-adapted, high-fidelity clearing strategies, deep multi-target labeling systems, and compatible 3D imaging platforms has become an urgent priority and a key technical breakthrough needed to advance both physiological and pathophysiological studies of the liver.

Removed in the revised version: An ideal three-dimensional (3D) visualization technique for the liver should enable clear delineation of the hierarchical architecture and structural alterations of the vascular systems at a macroscopic level, while simultaneously allowing the observation of terminal vascular structures coordinating with liver lobules at a single- cell resolution. Considering the unique anatomical and cellular characteristics of the liver, our research team has developed an advanced 3D visualization strategy tailored for the hepatic multi-ductal system, building upon existing imaging technologies. The core innovations of this strategy include:(1) Establishment of an efficient liver clearing protocol (liver-CUBIC): By systematically optimizing sample processing procedures and reagent compositions within the CUBIC tissue clearing framework, we achieved a 63.89% improvement in clearing efficiency and a 20.12% increase in tissue transparency.(2) Selection and application of multicolor metal compound nanoparticles (MCNPs): These nanoparticles exhibit excellent water resistance and compatibility with the Liver-CUBIC system. Their multicolor properties enable stable, simultaneous 3D labeling of the portal vein, hepatic artery, bile ducts, and central vein in mouse livers. Moreover, the MCNP series possesses unique excitation/emission spectra, making them suitable for multimodal imaging applications including brightfield, extended depth of field (EDF), and fluorescence microscopy. Their high spatial resolution supports the visualization of fine vascular terminal branches and subcellular structures.(3) Development of a multimodal immunolabeling system: This 3D visualization strategy is compatible with both DAB chromogenic and TSA-based multiplex immunofluorescence staining, enabling simultaneous labeling of various cell types and molecular antigens alongside vascular system visualization. This allows for the generation of high signal-to-noise ratio, multi-target, subcellular-resolution 3D structural maps of the liver. In summary, this newly established 3D visualization platform for hepatic multi-ductal systems provides both macroscopic and microscopic insights, offering a powerful tool for investigating the spatial organization, functional coordination, and pathological remodeling mechanisms of the liver microenvironment.

Removed in the revised version: Based on this platform, we successfully applied the proposed strategy to investigate the microanatomical architecture of the hepatic vascular system in mice, with a particular focus on elucidating how the complex intraparenchymal vascular networks are spatially organized and coordinate with liver lobule units at single-cell resolution. Through multimodal microscopic visualization, multicolor antigen-specific labeling, and targeted gene expression analysis, we identified a previously unrecognized functional structure termed the Periportal Lamellar Complex (PLC). This novel structure is distributed periodically along the portal vein axis and is closely associated with terminal bile duct branches and autonomic nerve plexuses, forming a specialized spatial interaction unit. Single-cell transcriptomic analysis combined with multiplex 3D immunostaining confirmed that the PLC possesses distinct spatial localization, morphological features, and molecular marker profiles within the murine liver. Its unique gene expression signatures and intimate interactions with multiple terminal vascular branches suggest a cooperative mechanism underlying hepatic microanatomical organization.

Earlier version: This discovery not only validates the superior capability of our multi-ductal 3D visualization platform in resolving microscale hepatic vascular interactions but also provides a technical foundation for establishing a new paradigm in hepatic vascular research. By enabling full-scale visualization from the organ to the vascular network and down to the single-cell level, this platform offers a powerful methodological tool for uncovering regulatory mechanisms of hepatic physiology and pathological remodeling.

Revised version: In summary, this liver-specific multicolor 3D visualization system enables comprehensive mapping of hepatic ductal-vascular interactions from the organ to the single-cell level. It not only demonstrates the power of the platform in resolving complex hepatic microstructures but also establishes a new technical foundation for investigating hepatic physiology and pathological remodeling.

Added in the revised version: Results

Earlier version: H2O₂-Enhanced CUBIC Clearing Coupled with Multicolor Nanoprobes for High- Resolution 3D Mapping of Liver Ductal -Vascular network

Revised version: H2O2-enhanced CUBIC clearing with multicolor nanoparticles enables high-resolution 3D mapping of the hepatic ductal–vascular network

Earlier version: The CUBIC (Clear, Unobstructed Brain/Body Imaging Cocktails and Computational analysis) technique has been widely applied in the tissue clearing and 3D imaging of organs such as the brain, kidney, and heart due to its excellent compatibility with various fluorescent dyes and immunolabeling protocols[15, 16]. However, the application of CUBIC in liver tissue remains challenging because of the liver’s high heme content and densely packed cellular architecture, which leads to pronounced light scattering and markedly limits the clearing efficacy. In addition, the conventional CUBIC protocol requires up to 9 days for liver tissue processing, severely restricting experimental throughput (Figure 1A, red schematic).To address these limitations, Laura M. Molina et al. incorporated a hydrogen peroxide (H2O2) bleaching step into the CLARITY protocol, improving liver transparency but failing to reduce the processing time; clearing efficiency also remained suboptimal [14] In this study, we developed an optimized clearing strategy termed Liver-CUBIC, featuring two key modifications:(1) Integration of an H2O2-CUBIC hybrid system, incorporating an H2O2 bleaching step into the CUBIC workflow. This dual-action mechanism effectively removes pigments such as heme and lipofuscin through oxidative reactions, while partially clearing lipid components to reduce light scattering.(2) Optimization of urea concentration, increasing from 25% to 40%, enhancing protein denaturation to minimize light scattering and improve tissue permeability, thus promoting more uniform reagent distribution. Experimental results confirmed that the Liver-CUBIC protocol significantly improved liver clearing efficiency (Figure S1A). Specifically, initial PBS/PFA perfusion efficiently removed residual heme (Figure S1A, panel 2); subsequent H2O2 bleaching eliminated pigment deposits (Figure S1A, panel 3); and final 40% urea treatment achieved refractive index homogenization (Figure S1A, panel 4).Using this optimized protocol, the total clearing time was reduced from 9 days with conventional CUBIC (25% urea, Figure 1A, red schematic) to 3.25 days with 40% urea + H2O2 treatment (Figure 1A, green schematic), representing a 63.89% increase in efficiency. Importantly, no significant tissue expansion was observed following Liver-CUBIC treatment (Figures 1B and 1D).Transmittance measurements revealed that the Liver-CUBIC protocol (40% urea + H2O2) markedly outperformed conventional CUBIC, with an average transmittance improvement of 20.12% (P < 0.0001; 95% CI: 19.14–21.09) (Figure 1E). Additionally, in 1 mm-thick liver sections, the clearing time was reduced to just 20 minutes (Figure S1B), while real-time imaging demonstrated complete clearing of 200 μm-thick slices within 60 seconds (Figure S1C; Supplementary Video 1), fully meeting the demands for rapid live imaging.

Revised version: The CUBIC (Clear, Unobstructed Brain/Body Imaging Cocktails and Computational analysis) technique has been widely applied in the tissue clearing and 3D imaging of organs such as the brain, kidney, and heart due to its excellent compatibility with various fluorescent dyes and immunolabeling protocols (Susaki et al., 2014; Tainaka et al., 2014). However, the application of CUBIC in liver tissue remains challenging because of the liver’s high heme content and densely packed cellular architecture, which leads to pronounced light scattering and markedly limits the clearing efficacy. Furthermore, the conventional CUBIC protocol requires up to 9 days for liver tissue processing, severely restricting experimental throughput (Figure 1A, red schematic). To address these limitations, Laura M. Molina et al. incorporated a hydrogen peroxide (H2O2) bleaching step into the CLARITY protocol, improving liver transparency but failing to reduce the processing time (Molina et al., 2022). In this study, we developed an optimized clearing strategy termed Liver-CUBIC, featuring two key modifications: (1) Integration of an H2O2-CUBIC hybrid system, incorporating an H2O2 bleaching step into the CUBIC workflow. This technical improvement effectively removes pigments such as heme and lipofuscin through oxidative reactions, while partially clearing lipid components to reduce light scattering; (2) Optimization of urea concentration, increasing from 25% to 40%, enhancing protein denaturation to minimize light scattering and improve tissue permeability, thus promoting more uniform reagent distribution. Experimental results confirmed that the Liver-CUBIC protocol significantly improved liver clearing efficiency (Figure 1—figure supplement 1A). Specifically, initial PBS/PFA perfusion efficiently removed residual heme; subsequent H2O2 bleaching eliminated pigment deposits; and final 40% urea treatment achieved refractive index homogenization (Figure 1—figure supplement 1A). Using this optimized protocol, the total clearing time was reduced from 9 days with conventional CUBIC (Figure 1A, red schematic) to 3.25 days with 40% urea + H2O2 treatment (Figure 1A, green schematic), representing a 63.89% increase in efficiency. Importantly, no significant tissue expansion was observed following Liver-CUBIC treatment (Figure 1B-D). Transmittance measurements revealed that the Liver-CUBIC protocol (40% urea + H2O2) markedly outperformed conventional CUBIC, with an average transmittance improvement of 20.12% (P < 0.0001; 95% CI: 19.14–21.09) (Figure 1E). Additionally, in 1 mm-thick liver sections, the clearing time was reduced to just 20 minutes (Figure 1—figure supplement 1B), while real-time imaging demonstrated complete clearing of 200 μm-thick slices within 60 seconds (Figure 1—figure supplement 1C, and Figure 1—video 1), fully meeting the demands for rapid live imaging.

Removed in the revised version: Metal compound nanoparticles (MCNPs) exhibit excellent brightness, tunable multicolor properties, and strong resistance to H2O2 bleaching, maintaining color stability even after Liver-CUBIC clearing. Utilizing four distinct MCNP colors — pink, green, black, and yellow — we developed a four-channel vascular labeling strategy (Figures 1F and S1D). The hepatic artery was perfused with yellow MCNPs via the left ventricle, bile ducts were retrogradely injected with green MCNPs through the extrahepatic duct, portal veins were directly filled with pink MCNPs, and central veins were labeled using black MCNPs via the inferior vena cava Using this method, simultaneous four-color labeling of the intrahepatic ductal-vascular system can be achieved (Figures S1E–S1G), while single- or dual-color labeling allows for higher- resolution visualization of fine microscale structures. High-resolution 3D imaging revealed that MCNP-Green-labeled large intrahepatic bile ducts exhibit a one-to-one parallel topology with adjacent portal vein branches. At the periphery, the terminal biliary network forms a classic three-dimensional plexus, characterized by helical configurations and polygonal extensions projecting into the hepatic lobules (Figure 1G).Meanwhile, MCNP-Pink-labeled portal veins displayed a typical tree-like branching pattern in the central regions of hepatic lobes, but in the marginal areas, they formed an unusually dense, radial micro-branching pattern (Figure 1H). Notably, the terminal micro-branches of the portal vein appeared not only around secondary branches but also intermittently along the main trunks, challenging the conventional notion of strictly hierarchical vascular bifurcation. For MCNP-Black-labeled central veins, imaging revealed sparse, circular or oval fenestrae measuring 50–100 nm in diameter distributed across the surface of the central vein trunks (Figure 1I). Previous scanning electron microscopy studies have described similar structures, which are thought to facilitate the return of hepatocyte-synthesized metabolic products (such as lipoproteins and cytokines) into the circulation, assist immune cell trafficking from the liver to the systemic circulation, and help maintain hepatic sinusoid–interstitium osmotic pressure balance[18]. Additionally, dual-label imaging using MCNP-Pink and MCNP-Green demonstrated that the hepatic artery runs parallel to the portal vein trunk, with the hepatic artery displaying a notably smaller diameter compared to its adjacent portal vein. At the periphery, hepatic arterial branches formed a three-dimensional entwined network that closely contacted the portal vein micro-branches (Figure 1J).

Earlier version: H2O2-Enhanced CUBIC Clearing Coupled with Multicolor Nanoprobes for High-Resolution Mapping of Liver Vasculature.

Revised version: H2O2-enhanced CUBIC clearing with multicolor nanoparticles enables high-resolution 3D mapping of the hepatic ductal–vascular network.

Earlier version: (A) Comparison of processing times for four different liver clearing protocols: conventional CUBIC (25% urea, red), high-concentration urea (40% urea, yellow), oxidation treatment alone (25% urea + 4.5% H2O2, blue), and the optimized Liver-CUBIC (40% urea + 4.5% H2O2, green). The optimized Liver-CUBIC protocol significantly reduced clearing time (n = 6 per group). (B) Schematic illustration of mouse liver lobes, defining left lateral lobe, left medial lobe, right medial lobe, right lateral lobe, caudate lobe, and quadrate lobe according to reference [19] . (C) Brightfield images of whole mouse livers (8–9 weeks old, male) after processing with the four clearing protocols. Grid size: 1.6 mm × 1.6 mm. Scale bar: 8 mm. (D) Quantitative analysis of tissue volume changes following each clearing protocol (n = 6 per group). No statistically significant difference was observed between the 25% urea group and the 40% urea + H2O2 group. Statistical test: unpaired two-tailed t-test. PT: pre-treatment. (E) Transmission spectra (400–900 nm) of 1 mm-thick mouse liver samples after each clearing protocol (n = 5 per group). (F) Schematic diagram of the four-channel ductal-vascular labeling strategy: yellow MCNPs injected via the left ventricle to label hepatic arteries, green MCNPs via retrograde common bile duct injection to label biliary ducts, pink MCNPs via the portal vein trunk to label portal veins, and black MCNPs via the inferior vena cava to label hepatic veins. (G) Three-dimensional fine structures of the biliary tree in the central and peripheral regions of the mouse liver, labeled with MCNP-Green. Images are shown at magnifications of 100×, 200×, and 400×. The rightmost panel presents a high-magnification view of the area outlined by the blue box. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm. The arrow indicates the terminal ductal structures with a polygonal shape. (H) Three-dimensional reconstruction of the portal venous system labeled with MCNP-Pink, showing branching features in central and peripheral zones. Images captured at 100×, 200×, and 400× magnifications; the rightmost image is a higher-magnification view of the blue-boxed region. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm. (I) Three-dimensional structural details of the central vein labeled with MCNP-Black. Images captured at 100×, 300×, and 400× magnifications; the rightmost image is a higher-magnification view of the blue-boxed region. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm. (J) Three-dimensional structural details of the hepatic artery labeled with MCNP-Pink. Images captured at 100×, 200×, and 400× magnifications; the rightmost image is a higher-magnification view of the blue-boxed region. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm.

Revised version: (A) Comparison of processing times for four different liver clearing protocols: conventional CUBIC (25% urea, red), high-concentration urea (40% urea, yellow), oxidation treatment alone (25% urea + 4.5% H2O2, blue), and the optimized Liver-CUBIC (40% urea + 4.5% H2O2, green). The optimized Liver-CUBIC protocol significantly reduced clearing time (n = 6 per group). (B) Schematic illustration of mouse liver lobes, defining left lateral lobe, left medial lobe, right medial lobe, right lateral lobe, caudate lobe, and quadrate lobe according to reference (Y. Zhang & DeBosch, 2022). (C) Brightfield images of whole mouse livers (8–9 weeks old, male) after processing with the four clearing protocols. Grid size: 1.6 mm × 1.6 mm. Scale bar: 8 mm. (D) Quantitative analysis of tissue volume changes following each clearing protocol (n = 6 per group). No statistically significant difference was observed between the 25% urea group and the 40% urea + PT group. Statistical test: unpaired two-tailed t-test. PT: pre-treatment. (E) Transmission spectra (400–900 nm) of 1 mm-thick mouse liver samples after each clearing protocol (n = 5 per group). (F) Schematic diagram of the four-channel ductal-vascular labeling strategy: yellow MCNPs injected via the left ventricle to label hepatic arteries, green MCNPs via retrograde common bile duct injection to label biliary ducts, pink MCNPs via the portal vein trunk to label portal veins, and black MCNPs via the inferior vena cava to label hepatic veins. (G) Three-dimensional fine structures of the biliary tree in the central and peripheral regions of the mouse liver, labeled with MCNP-Green. Images are shown at magnifications of 100 ×, 200 ×, and 400 ×. The rightmost panel presents a high-magnification view of the area outlined by the blue box. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm. The arrow indicates the terminal ductal structures with a polygonal shape. (H) Three-dimensional reconstruction of the portal venous system labeled with MCNP-Pink, showing branching features in central and peripheral zones. Images captured at 100 ×, 200 ×, and 400 × magnifications; the rightmost image is a higher-magnification view of the blue-boxed region. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm. (I) Three-dimensional structural details of the central vein labeled with MCNP-Black. MCNP-Yellow, MCNP-Pink, and MCNP-Green were used to label the portal vein, hepatic artery, and bile duct, respectively. Images captured at 100×, 300×, and 400× magnifications; the rightmost image is a higher-magnification view of the blue-boxed region. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm. (J) Three-dimensional structural details of the hepatic artery labeled with MCNP-Pink. Images captured at 100 ×, 200 ×, and 400 × magnifications; the rightmost image is a higher-magnification view of the blue-boxed region. Scale bars: 400 μm, 200 μm, 100 μm, and 20 μm.

Added in the revised version: Metal compound nanoparticles (MCNPs) exhibit excellent brightness, tunable multicolor properties, and strong resistance to H2O2 bleaching, maintaining color stability even after Liver-CUBIC clearing. Utilizing four distinct MCNP colors — pink, green, black, and yellow — we developed a four-channel vascular labeling strategy (Figure 1F and Figure 1—figure supplement 1D). The hepatic artery was perfused with yellow MCNPs via the left ventricle, bile ducts were retrogradely injected with green MCNPs through the extrahepatic duct, portal veins were directly filled with pink MCNPs, and central veins were labeled using black MCNPs via the inferior vena cava. Using this method, simultaneous four-color labeling of the intrahepatic ductal-vascular system can be achieved (Figure 1—figure supplement 1E-G), while single- or dual-color labeling allows for higher-resolution visualization of fine microscale structures. High-resolution 3D imaging revealed that MCNP-Green-labeled large intrahepatic bile ducts exhibit a one-to-one parallel topology with adjacent portal vein branches. At the periphery, the terminal biliary network forms a classic three-dimensional plexus, characterized by helical configurations and polygonal extensions projecting into the hepatic lobules (Figure 1G). Meanwhile, MCNP-Pink-labeled portal veins displayed a typical tree-like branching pattern in the central regions of hepatic lobes, but in the marginal areas, they formed an unusually dense, radial micro-branching pattern (Figure 1H). Notably, the terminal micro-branches of the portal vein appeared not only around secondary branches but also intermittently along the main trunks, challenging the conventional notion of strictly hierarchical vascular bifurcation. For MCNP-Black-labeled central veins, imaging revealed sparse, circular or oval fenestrae measuring 5–10 μm in diameter distributed across the surface of the central vein trunks (Figure 1I). Previous scanning electron microscopy studies have described similar structures, which are thought to facilitate the return of hepatocyte-synthesized metabolic products (such as lipoproteins and cytokines) into the circulation, assist immune cell trafficking from the liver to the systemic circulation, and help maintain hepatic sinusoid–interstitium osmotic pressure balance (Mak & Shin, 2021). Additionally, dual-label imaging using MCNP-Pink and MCNP-Green demonstrated that the hepatic artery runs parallel to the portal vein trunk, with the hepatic artery displaying a notably smaller diameter compared to its adjacent portal vein. At the periphery, hepatic arterial branches formed a three-dimensional entwined network that closely contacted the portal vein micro-branches (Figure 1J).

In summary, combining Liver-CUBIC with MCNPs enabled highly efficient, simultaneous mapping of intrahepatic bile ducts, portal veins, central veins, and hepatic arteries in mice at cellular resolution, revealing their spatial organization and topological relationships.

The periportal lamellar complex (PLC) serves as a low-permeability gateway bridging portal veins to hepatic lobules

Added in the revised version: In this study, scanning electron microscopy revealed that MCNP dye is a nanoparticle-based labeling agent with a diameter of approximately 150-200 nm (Figure 2A). Our experimental results showed that MCNPs were primarily distributed within the lumens of blood vessels and bile ducts through physical perfusion, without exhibiting specific molecular binding capacity. Under our perfusion conditions, the dye signals remained confined to the luminal regions, clearly delineating the morphology of vascular and biliary structures without diffusing into the surrounding parenchyma (Figure 2—figure supplement 1A-H). However, due to the increased imaging depth of liver tissue after optical clearing, nuclear stains such as DAPI often generate highly dense fluorescence signals that obscure fine vascular and perivascular structures (Figure 2—figure supplement 1I, J). Thus, our imaging analyses primarily focused on the high-resolution visualization of the spatial architecture of the vascular and biliary systems.

Removed in the revised version: In this study, scanning electron microscopy revealed that MCNP dye is a nanoparticle-based labeling agent with a diameter of approximately 100 nm (Figure 2A). We employed a dual-channel vascular labeling strategy to simultaneously visualize the portal vein and hepatic artery systems in the mouse liver. Specifically, MCNP-Green was employed to label the portal venous system by injection from portal vein, while MCNP- Pink marked the hepatic arterial system by injection extrahepatic duct(Figure 2B).

The periportal lamellar complex (PLC) serves as a low-permeability gateway bridging portal veins to hepatic lobules.

Earlier version: (A) Scanning electron microscopy revealed that the Metal Compound Nanoparticles (MCNP) dyes consisted of aggregates of particles approximately 100 nm in length. (B) Dual-channel vascular typing labeling scheme (Dual-channel NanoFluor™). hepatic arteries labeled with pink fluorescent nanoparticlesportal veins labeled with green fluorescent nanoparticles. (C) Extended depth-of-field imaging (tissue thickness: 200 μm) reveals fine three-dimensional structural details of the hepatic artery and portal vein. The right panel shows magnified lateral and top views of the orange-boxed region, highlighting dense terminal branches and the close spatial proximity between the two vascular systems. Green arrows indicate portal veins, and pink arrows indicate hepatic arteries. (D) Magnified image showing periodic alignment of PLC structures (orange arrows) along the adventitial layer of the portal vein trunk (green) in the direction of vascular flow. Orange dashed lines delineate the boundaries of a classical hepatic lobule. Scale bar: 200 μm. (E) Three-dimensional confocal imaging of hepatic arteries (pink dye, excitation/emission: 561/640 nm) and portal veins (green dye, excitation/emission: 495/519 nm). The right panel presents a high-resolution view with hepatic arteries rendered in blue and the portal vein surface filled in white; arrows indicate points of interaction between PLC structures and hepatic arteries. Scale bar: 200 μm. (F) High-magnification confocal imaging further depicting micro-branches of portal veins and hepatic arteries, with terminal branches intertwining in a coiled distribution. Scale bar: 20 μm. (G, H) Single-channel confocal images showing the distribution of MCNP-Green-labeled portal veins (G) and MCNP- Pink-labeled hepatic arteries (H) within PLC regions. White arrows indicate the paths of fluorescence intensity profile measurements, with arrowheads denoting the direction of line scans. Scale bar: 50 μm. (I) Schematic illustration of fluorescence intensity profile measurements across PLC structures. The midpoint of the portal vein (green) intensity profile corresponds to the junction between the PLC and the outer wall of the portal vein, while the midpoint of the hepatic artery (pink) intensity profile aligns with the terminal edge of the PLC adjacent to liver sinusoids. Arrows indicate scan directions. (G) Fluorescence intensity profile plots. The X-axis represents the scan distance (0–400 μm), and the Y-axis represents fluorescence intensity. Both portal vein (green) and hepatic artery (pink) signals showed significant increases within the PLC region. The regions were defined as follows: 0–100 μm, portal vein region; 100–250 μm, PLC region; 250–400 μm, liver sinusoid region (n = 5 per group). (K) Analytical workflow for characterizing PLC structures: the primary portal vein trunk carrying PLC structures was selected as the reference axis for hepatic lobule boundaries, combined with its two adjacent secondary branches to define a classical hexagonal lobule computational unit. The diameter of the primary portal vein trunk (dPV) and the area of PLC structures were quantified using the extended depth-of-field imaging system. (L) Distribution of PLC areas along primary portal vein trunks with diameters ranging from 63.45 to 321.42 μm. Each value represents the PLC area associated with a portal vein of corresponding diameter (n = 19).

Revised version: (A) Scanning electron microscopy revealed that the Metal Compound Nanoparticles (MCNP) dyes consisted of aggregates of particles approximately 100-200 nm in length. (B) Dual-channel vascular typing labeling scheme (Dual-channel NanoFluor™). hepatic arteries labeled with pink fluorescent nanoparticles portal veins labeled with green fluorescent nanoparticles. (C) EDOF imaging (tissue thickness: 200 μm) reveals fine three-dimensional structural details of the hepatic artery and portal vein. The right panel shows magnified lateral and top views of the orange-boxed region, highlighting dense terminal branches and the close spatial proximity between the two vascular systems. Green arrows indicate portal veins, and pink arrows indicate hepatic arteries. (D) Magnified image showing periodic alignment of PLC structures (orange arrows) along the adventitial layer of the portal vein trunk (green) in the direction of vascular flow. Orange dashed lines delineate the boundaries of a classical hepatic lobule. Scale bar: 200 μm. (E) Three-dimensional confocal imaging of hepatic arteries (pink dye, excitation/emission: 561/640 nm) and portal veins (green dye, excitation/emission: 495/519 nm). The right panel presents a high-resolution view with hepatic arteries rendered in blue and the portal vein surface filled in white. Scale bar: 200 μm. (F) High-magnification confocal imaging further depicting micro-branches of portal veins and hepatic arteries, with terminal branches intertwining in a coiled distribution. arrows indicate points of interaction between PLC structures and hepatic arteries. Scale bar: 20 μm. (G, H) Single-channel confocal images showing the distribution of MCNP-Green-labeled portal veins (G) and MCNP-Pink-labeled hepatic arteries (H) within PLC regions. White arrows indicate the paths of fluorescence intensity profile measurements, with arrowheads denoting the direction of line scans. Scale bar: 50 μm. (I) Schematic illustration of fluorescence intensity profile measurements across PLC structures. The midpoint of the portal vein (green) intensity profile corresponds to the junction between the PLC and the outer wall of the portal vein, while the midpoint of the hepatic artery (pink) intensity profile aligns with the terminal edge of the PLC adjacent to liver sinusoids. Arrows indicate scan directions. (J) Fluorescence intensity profile plots. The X-axis represents the scan distance (0–400 μm), and the Y-axis represents fluorescence intensity. Both portal vein (green) and hepatic artery (pink) signals showed significant increases within the PLC region. The regions were defined as follows: 0–100 μm, portal vein region; 100–250 μm, PLC region; 250–400 μm, liver sinusoid region (n = 5 per group). (K) Analytical workflow for characterizing PLC structures: the primary portal vein trunk carrying PLC structures was selected as the reference axis for hepatic lobule boundaries, combined with its two adjacent secondary branches to define a classical hexagonal lobule computational unit. The diameter of the primary portal vein trunk (dPV) and the area of PLC structures were quantified using the extended depth-of-field imaging system. (L) Distribution of PLC areas along primary portal vein trunks with diameters ranging from 63.45 to 321.42 μm. Each value represents the PLC area associated with a portal vein of corresponding diameter (n = 19).

Added in the revised version: We employed a dual-channel vascular labeling strategy to simultaneously visualize the portal vein and hepatic artery systems in the mouse liver. Specifically, MCNP-Green was employed to label the portal venous system by injection from portal vein, while MCNP-Pink marked the hepatic arterial system by injection via left ventricular (Figure 2B).

Earlier version: Extended depth-of-field imaging revealed that MCNP-Green specifically labels a distinct vascular structural unit located in the adventitial layer of the portal vein trunk, which we designate as the Periportal Lamellar Complex (PLC) (Figures 2C and 2D, orange arrows). Morphologically, this structure consists of fine vascular branches distributed along the surface of the portal vein trunk. Its base is anchored to the portal vein, and it radiates outward in a lobular pattern centered on the portal vein, giving rise to multiple terminal branches that directly connect to the liver sinusoids (Figure 2C).

Revised version: Extended depth-of-field (EDOF) imaging revealed that MCNP-Green specifically labels a distinct vascular structural unit located in the adventitial layer of the portal vein trunk, which we designate as the Periportal Lamellar Complex (PLC) (Figure 2C, D, orange arrows). Morphologically, this structure consists of fine vascular branches distributed along the surface of the portal vein trunk. Its base is anchored to the portal vein, and it radiates outward in a lobular pattern centered on the portal vein, giving rise to multiple terminal branches that directly connect to the liver sinusoids (Figure 2C). The PLC structures align in a serial arrangement along the portal vein trunk (Figure 2D).

Removed in the revised version: The PLC structures align in a serial arrangement along the portal vein trunk (Figure 2D)

Earlier version: Further validation by confocal fluorescence imaging confirmed that following portal vein perfusion with MCNP-Green, fluorescence signals are highly enriched within the PLC. Quantitative analysis showed that the fluorescence intensity within the PLC was significantly higher than that observed in the portal vein lumen and liver sinusoid regions, with minimal dye diffusion at the PLC periphery, indicating predominant localization within the main PLC structure (Figures 2G–2J). In contrast, MCNP-Pink signals, introduced via hepatic artery perfusion, predominantly localized within the PLC and its margins without crossing into adjacent PLC units. Due to differences in perfusion pathways, the pink dye demonstrated a greater tendency to extend from the PLC towards the lobular parenchyma compared to the green dye. This phenomenon likely reflects the relatively low permeability of the PLC region: green nanoparticles delivered through the portal vein are largely retained within the low-permeability PLC core, with limited peripheral diffusion, whereas the arterial pink dye, influenced by the pressure and flow characteristics of the hepatic artery system, effectively labels the peripheral PLC structures (Figures 2G–2J).

Revised version: Further validation by confocal fluorescence imaging confirmed that following portal vein perfusion with MCNP-Green, fluorescence signals are highly enriched within the PLC. Quantitative analysis showed that the fluorescence intensity within the PLC was significantly higher than that observed in the portal vein lumen and liver sinusoid regions, with minimal dye diffusion at the PLC periphery, indicating predominant localization within the main PLC structure (Figure 2E-J). In contrast, MCNP-Pink signals, introduced via hepatic artery perfusion, predominantly localized within the PLC and its margins without crossing into adjacent PLC units. Due to differences in perfusion pathways, the pink dye demonstrated a greater tendency to extend from the PLC towards the lobular parenchyma compared to the green dye. This phenomenon likely reflects the relatively low permeability of the PLC region: green nanoparticles delivered through the portal vein are largely retained within the low-permeability PLC core, with limited peripheral diffusion, whereas the arterial pink dye, influenced by the pressure and flow characteristics of the hepatic artery system, effectively labels the peripheral PLC structures (Figure 2G-J).

Earlier version: Using the Extended Depth-of-Field Multimodal Imaging System, we quantitatively analyzed the distribution of PLC structures in relation to portal vein diameter and anatomical location. By defining classical hexagonal hepatic lobule units through the primary portal vein trunk containing PLC structures and its two adjacent secondary branches (Figure 2H), quantitative analysis revealed that the length of the primary portal vein trunk ranged from 269.85 μm to 1513.67 μm (n = 19), with a median length of 769.89 μm (IQR: 593.67 μm–987.77 μm). The number of PLC structures per trunk ranged from 1 to 4, with 68.42% (13/19) of portal veins containing two PLCs (Figure S2A). Notably, neither the number nor the area of PLCs showed significant correlation with the length of the portal vein trunk (Figures S2A and S2B), indicating that morphological parameters alone (such as trunk length) may be insufficient to define hepatic lobular units, especially given the substantial differences in blood transport capacity among portal veins of varying diameters. The diameters of the primary portal veins bearing PLC structures ranged from 63.45 μm to 321.42 μm, and PLC area showed a significant positive correlation with portal vein diameter (Figure 2L), whereas PLC number did not correlate with diameter (Figure S2C). This diameter–area relationship may reflect a functional adaptation of the portal venous system in metabolically active liver regions: larger-diameter portal veins accommodate higher blood flow and consequently experience greater shear stress on the vessel wall. The formation of PLC structures in the adventitial layer may participate in local blood flow regulation, maintenance of microenvironmental homeostasis, and vascular–stem cell interactions.

Revised version: To further characterize PLC distribution, we used the EDOF imaging system to quantitatively analyze the distribution of PLC structures in relation to portal vein diameter and anatomical location. By defining classical hexagonal hepatic lobule units through the primary portal vein trunk containing PLC structures and its two adjacent secondary branches (Figure 2K), quantitative analysis revealed that the length of the primary portal vein trunk ranged from 269.85 μm to 1513.67 μm (n = 19), with a median length of 769.89 μm (IQR: 593.67 μm – 987.77 μm), the number of PLC structures per trunk ranged from 1 to 4, with 68.42% (13/19) of portal veins containing two PLCs (Figure 2—figure supplement 1K). Notably, neither the number nor the area of PLCs showed significant correlation with the length of the portal vein trunk (Figure 2—figure supplement 1K, L), indicating that morphological parameters alone (such as trunk length) may be insufficient to define hepatic lobular units, especially given the substantial differences in blood transport capacity among portal veins of varying diameters. The diameters of the primary portal veins bearing PLC structures ranged from 63.45 μm to 321.42 μm, and PLC area showed a significant positive correlation with portal vein diameter (Figure 2L), whereas PLC number did not correlate with diameter (Figure 2—figure supplement 1M). This diameter–area relationship may reflect a functional adaptation of the portal venous system in metabolically active liver regions: larger-diameter portal veins accommodate higher blood flow and consequently experience greater shear stress on the vessel wall. The formation of PLC structures in the adventitial layer may participate in local blood flow regulation, maintenance of microenvironmental homeostasis.

These findings suggest that the Periportal Lamellar Complex (PLC) is not only a morphologically and spatially distinct, low-permeability vascular unit surrounding the portal vein, but also likely serves as a critical nexus connecting the portal vein, hepatic artery, and liver sinusoids. Thus, the PLC constitutes a key node within the interactive vascular network of the mouse liver.

Spatial juxtaposition of the periportal lamellar complex with canals of Hering at the portal venous interface

Earlier version: To further investigate the potential relationship between the PLC and bile delivery regulation within the hepatic lobule, we employed a dual-channel vascular labeling strategy to simultaneously visualize the mouse liver portal vein and bile duct systems.

Revised version: To further investigate the potential relationship between the PLC and bile delivery regulation within the hepatic lobule, we employed a dual-channel vascular-biliary labeling strategy to simultaneously visualize the mouse liver portal vein and bile duct systems.

Earlier version: We combined MCNP-Pink labeling of the portal vein and its associated PLC structures with three-dimensional DAB immunohistochemical staining for CK19 to visualize the bile duct epithelial architecture. Extended depth-of-field three-dimensional bright-field imaging revealed a strict 1:1 anatomical association between the primary portal vein trunk (diameter 280 ± 32 μm) and the first-order bile duct (diameter 69 ± 8 μm) (Figures 3A and S3A).From multiview imaging showed that second- order bile ducts are distributed along the portal vein trunk and branch directionally toward the PLC. Their terminal portions, namely the canals of Hering, form the distal bile duct network responsible for bile collection, and are located at the interface between hepatic lobules, rather than being evenly distributed around the portal vein.

Revised version: We combined MCNP-Pink labeling of the portal vein and its associated PLC structures with three-dimensional DAB immunohistochemical staining for CK19 to visualize the bile duct epithelial architecture. EDOF imaging revealed a consistent 1:1 anatomical association between the primary portal vein trunk (diameter 280 ± 32 μm) and the first-order bile duct (diameter 69 ± 8 μm) (Figure 3A, Figure 3—figure supplement 1A). Multiview imaging showed that second-order bile ducts are distributed along the portal vein trunk and branch directionally toward the PLC. The terminal portions, namely the canals of Hering, form the distal bile duct network responsible for bile collection and are located at the interface between hepatic lobules, rather than being evenly distributed around the portal vein. These terminal ducts are specifically concentrated within the PLC surrounding the portal vein, displaying marked spatial colocalization (Figure 3A). This finding suggests a potential anatomical spatial association between the terminal canals of Hering and the periportal PLC structures. Furthermore, based on MCNP-Green labeling of the biliary tract, combined with 3D DAB-CK19 staining, both EDOF imaging and 3D confocal fluorescence imaging consistently confirmed that the terminal regions of the biliary network—responsible for collecting bile from the hepatic parenchyma—are primarily located adjacent to the PLC (Figure 3B, C, Figure 3—figure supplement 1B).

Earlier version: Spatial Juxtaposition of the Periportal Lamellar Complex with Canals of Hering at the Portal Venous Interface

Revised version: Spatial Juxtaposition of the Periportal Lamellar Complex with Canals of Hering at the Portal Venous Interface Spatial juxtaposition of the periportal lamellar complex with canals of Hering at the portal venous interface.

Earlier version: (A) MCNP-Pink labeling of portal veins combined with three-dimensional DAB immunohistochemistry for CK19 (brown) to visualize bile duct epithelial cells. Top and lateral views highlight the PLC–bile duct interaction sites (cyan arrows). Scale bar: 50 μm. (B) MCNP-Green labeling of bile ducts combined with three-dimensional DAB immunohistochemistry for CK19 (brown). Scale bar: 40 μm. (C) MCNP-Green labeling of bile ducts combined with three-dimensional TSA immunofluorescence for CK19 (red), displaying detailed structures at the interface between dye-labeled ducts and immunostained bile duct terminals. The arrows indicate the terminal structures of the bile ducts. Scale bar: 10 μm. (D) Three-dimensional TSA multiplex immunofluorescence staining for ZO-1 (red, marking bile canaliculi networks), HNF4α (gray, marking hepatocyte nuclei), and CK19 (green, marking bile ducts). Scale bar: 20 μm. The right panel illustrates the spatial relationship between ZO-1-labeled bile canaliculi and CK19-labeled bile duct terminals, with arrows indicating the terminal positions of bile ducts. (E) Extended depth-of-field imaging of the whole liver showing high-pressure perfusion of green fluorescent nanoparticles into the bile duct, yellow nanoparticles labeling the portal vein. Red circles indicate sites of green dye leakage localized to the PLC regions. (F) High-pressure retrograde perfusion of red fluorescent nanoparticles into the bile duct, combined with three-dimensional TSA immunofluorescence for CK19 (green). Arrows indicate sites of dye leakage at the PLC region. Scale bar: 50 μm. (G) Schematic diagram illustrating the leakage sites of bile duct-perfused dye following high-pressure injection. Green dashed boxes represent the positions of free bile duct terminal epithelial cells at leakage sites.

Revised version: (A) MCNP-Pink labeling of portal veins combined with three-dimensional DAB immunohistochemistry for CK19 (brown) to visualize bile duct epithelial cells. Top and lateral views highlight the PLC–bile duct interaction sites (cyan arrows). Scale bar: 50 μm. (B) MCNP-Green labeling of bile ducts combined with three-dimensional DAB immunohistochemistry for CK19 (brown). Scale bar: 40 μm. (C) MCNP-Green labeling of bile ducts combined with three-dimensional TSA immunofluorescence for CK19 (red), displaying detailed structures at the interface between dye-labeled ducts and immunostained bile duct terminals. The arrows indicate the terminal structures of the bile ducts. Scale bar: 10 μm. (D) Three-dimensional TSA multiplex immunofluorescence staining for ZO-1 (red, marking bile canaliculi networks), HNF4α (gray, marking hepatocyte nuclei), and CK19 (green, marking bile ducts). Scale bar: 20 μm. The right panel illustrates the spatial relationship between ZO-1-labeled bile canaliculi and CK19-labeled bile duct terminals, with arrows indicating the terminal positions of bile ducts. (E) EDOF imaging of the whole liver showing high-pressure perfusion of green fluorescent nanoparticles into the bile duct and yellow nanoparticles labeling the portal vein. Pink signal (arrow) indicates hepatic artery labeling, which may appear slightly displaced toward the central vein due to minor misalignment during perfusion. Red circles indicate sites of green dye leakage localized to the PLC regions. (F) High-pressure retrograde perfusion of red fluorescent nanoparticles into the bile duct, combined with three-dimensional TSA immunofluorescence for CK19 (green). Arrows indicate sites of dye leakage at the PLC region. Scale bar: 50 μm. (G) Statistical analysis of the number of dye leakage foci shown in the left panel (H), comparing the PLC with the perisinusoidal regions adjacent to the portal vein. *** indicates a highly significant difference between groups (P < 0.0001). (H) Schematic diagram illustrating the leakage sites of bile duct-perfused dye following high-pressure injection. Green dashed boxes represent the positions of free bile duct terminal epithelial cells at leakage sites.

Removed in the revised version: Instead, these terminal ducts are specifically concentrated within the PLC surrounding the portal vein, displaying marked spatial colocalization (Figure 3A). This finding suggests a potential anatomical spatial association between the terminal canals of Hering and the periportal PLC structures. Furthermore, based on MCNP-Green labeling of the biliary tract, combined with 3D DAB-CK19 staining, both extended depth-of-field bright-field imaging and 3D confocal fluorescence imaging consistently confirmed that the terminal regions of the biliary network—responsible for collecting bile from the hepatic parenchyma—are primarily located adjacent to the PLC (Figures 3B, 3C, and S3B).

Earlier version: To delineate the microscopic architecture underlying bile flow from the hepatic parenchyma into the bile duct network, we employed three-dimensional TSA-based multiplex immunofluorescence to simultaneously label the tight junction protein ZO1, the hepatocyte nuclear marker HNF4A, and the bile duct epithelial marker CK19 (Figure 3D). Three-dimensional confocal imaging demonstrated that bile canaliculi on the hepatocyte surface formed polygonal networks demarcated by ZO1, while terminal bile duct epithelial cells also organized into luminal structures via ZO1, arranged in either single- or multicellular rings, thereby establishing an interface between the canalicular network and the bile duct system. However, detailed morphological analysis suggested that this connection does not occur directly between the bile canaliculi and the terminal bile duct epithelial cells. Instead, bile duct epithelial cells at the terminal ducts extended partially along the canalicular network without directly participating in the formation of the bile duct lumen. This finding indicates that these cells may remain in an immature state or possess alternative stem/progenitor cell-like properties [20, 21]. To further verify the precise interface between the terminal bile ducts and the bile canalicular network, we performed retrograde perfusion of MCNP-Green dye through the common bile duct while progressively increasing the perfusion pressure. When the dye pressure exceeded a defined threshold, dye leakage was observed (Figures 3E-3G, and S3C). Dual-channel three-dimensional confocal imaging combined with CK19 immunostaining revealed that the sites of dye leakage did not coincide with the CK19-positive terminal bile duct epithelium, but instead were predominantly localized within regions adjacent to the PLC structures (Figures 3E-3G, and S3C).This observation further supports our earlier hypothesis: first-order bile ducts distribute along the portal vein trunk, secondary bile ducts branch directionally toward the PLC regions, and terminal bile duct branches converge spatially adjacent to the PLC, where they collect bile drained from the hepatic lobules. These results suggest that the PLC may serve not only as a spatial positional cue guiding bile duct growth and branching but also as a regulatory node involved in coordinating bile drainage from the hepatic lobule into the biliary network.

Revised version: To delineate the microscopic architecture underlying bile flow from the hepatic parenchyma into the bile duct network, we employed three-dimensional TSA-based multiplex immunofluorescence to simultaneously label the tight junction protein ZO1, the hepatocyte nuclear marker HNF4A, and the bile duct epithelial marker CK19 (Figure 3D). Three-dimensional confocal imaging demonstrated that bile canaliculi on the hepatocyte surface formed polygonal networks demarcated by ZO1, while terminal bile duct epithelial cells also organized into luminal structures via ZO1, arranged in either single- or multicellular rings, thereby establishing an interface between the canalicular network and the bile duct system. However, detailed morphological analysis suggested that this connection does not occur directly between the bile canaliculi and the terminal bile duct epithelial cells. Instead, bile duct epithelial cells at the terminal ducts extended partially along the canalicular network without directly participating in the formation of the bile duct lumen (Figure 3D). This finding indicates that these cells may remain in an immature state or possess alternative stem/progenitor cell-like properties (Banales et al., 2019; de Jong, van den Heuvel, Wells, & Porte, 2021).

Added in the revised version: To further verify the precise interface between the terminal bile ducts and the bile canalicular network, we performed retrograde perfusion of MCNP-Green dye through the common bile duct while progressively increasing the perfusion pressure. When the dye pressure exceeded a defined threshold, dye leakage was observed (Figure 3E-H, Figure 3—figure supplement 1C). Quantitative analysis of ultra–deep-focus images demonstrated that the dye leakage sites were primarily located adjacent to the PLC structures, rather than in the perisinusoidal regions surrounding the portal vein (Figure 3G).

Added in the revised version: This observation further supports our earlier hypothesis: first-order bile ducts distribute along the portal vein trunk, secondary bile ducts branch directionally toward the PLC regions, and terminal bile duct branches converge spatially adjacent to the PLC, where they collect bile drained from the hepatic lobules. These results suggest that the PLC may serve not only as a spatial positional cue guiding bile duct growth and branching but also as a regulatory node involved in coordinating bile drainage from the hepatic lobule into the biliary network.

Earlier version: Single-Cell Transcriptomics Identifies CD34⁺Sca-1⁺ as a Novel Endothelial Signature of the Periportal Lamellar Complex with Hematopoietic Niche Potential

Revised version: Single-cell transcriptomics identifies CD34+Sca-1+ as a unique endothelial signature of the periportal lamellar complex with hematopoietic niche potential

Earlier version: To identify the transcriptional profiles of endothelial cells (ECs) of PLC, we re- analyzed publicly available single-cell RNA sequencing (scRNA-seq) data of liver endothelial cells isolated by fluorescence-activated cell sorting (FACS) from the livers of three healthy adult mice[22]. Unsupervised clustering identified a total of 10 distinct cell clusters (Figures 4A and 4B). Among these, several clusters contained contaminating non-endothelial cell populations expressing marker genes for hepatocytes, T cells, cholangiocytes, and macrophages (Figures 4A and 4B). These non- endothelial populations were excluded from subsequent analyses.Based on previously reported endothelial cell cluster-specific marker genes[23, 24], the remaining liver endothelial cells were classified into five distinct subpopulations: portal vein endothelial cells (Portal vein ECs), periportal liver sinusoidal endothelial cells (Periportal LSECs), midzonal liver sinusoidal endothelial cells (Midzonal LSECs), pericentral liver sinusoidal endothelial cells (Pericentral LSECs), and central vein endothelial cells (Central vein ECs) (Figures 4C, S4A and S4B).

Revised version: To investigate the molecular characteristics and potential functional roles of endothelial cells (ECs) within the PLC, we re-analyzed publicly available single-cell RNA sequencing (scRNA-seq) data of liver endothelial cells isolated by fluorescence-activated cell sorting (FACS) from three healthy adult mouse livers (Su et al., 2021). Unsupervised clustering identified a total of 10 distinct cell clusters (Figure 4A, B and Figure 4—source data 1). Several clusters contained contaminating non-endothelial populations expressing marker genes for hepatocytes, T cells, cholangiocytes, and macrophages (Figure 4A, B), which were excluded from subsequent analyses. After removing these contaminating clusters, we classified the remaining liver ECs into five subpopulations based on previously reported endothelial cell cluster-specific marker genes (Halpern et al., 2018; Inverso et al., 2021), the remaining liver ECs were classified into five subpopulations: portal vein ECs, periportal LSECs, midzonal LSECs, pericentral LSECs, and central vein ECs (Figure 4C, Figure 4—figure supplement 1A, B).

Earlier version: Single-Cell Transcriptomics Identifies CD34⁺Sca-1⁺ as a Novel Endothelial Signature of the Periportal Lamellar Complex with Hematopoietic Niche Potential

Revised version: Single-cell transcriptomics identifies CD34+Sca-1+ as a unique endothelial signature of the periportal lamellar complex with hematopoietic niche potential.

Earlier version: (A) UMAP projection of single-cell transcriptomes of liver endothelial cells from normal adult mice reveals 10 distinct cellular clusters. Each dot represents one cell. (B) Heatmap showing expression profiles of representative genes across the 10 clusters, including hepatic stellate cells, T cells, macrophages, hepatocytes, cholangiocytes, prtal vein ECs, periportal LSECs, midzonal LSECs, pericentral LSECs, and central vein ECs . Marker genes were selected based on average cluster expression. (C) Representative cluster-specific markers for five major liver endothelial subpopulations. (D) Spatial schematic illustrating the anatomical position of the Periportal Lamellar Complex (PLC), located exclusively between the portal vein ECs and periportal LSECs. (E) Venn diagram showing overlapping top 20 highly expressed genes between prtal vein ECs and periportal LSECs. (F) Multiplex immunofluorescence images showing CD31 (yellow), Sca-1 (green), and CD34 (red) expression in distinct hepatic endothelial zones, including prtal vein ECs, midzonal LSECs, and central vein ECs. The PLC structure is demarcated by dashed white arrows. Scale bar: 40 μm. (G) Gating strategy and classification of CD34⁺Sca-1⁺ double-positive versus double-negative endothelial subpopulations. (H) Volcano plot illustrating differentially expressed hematopoietic-associated genes in CD34⁺Sca-1⁺ cells. The x-axis shows log₂ fold-change (log₂FC); the y-axis shows −log₁₀ adjusted P value. Significantly upregulated genes are located in the upper right quadrant. (I) Gene Ontology enrichment analysis of genes upregulated in CD34⁺Sca-1⁺ ECs, highlighting functional categories such as stem cell development, differentiation, and somatic stem cell maintenance. (J) The schematic illustration depicts the spatial localization of Sca-1, a marker for mesenchymal or hematopoietic stem cells, and CD34⁺Sca-1⁺CD31⁺ endothelial cells as the trunk of the Periportal Lamellar Complex (PLC).

Revised version: (A) UMAP projection of single-cell transcriptomes of liver endothelial cells from normal adult mice reveals 10 distinct cellular clusters. Each dot represents one cell. (B) Heatmap showing expression profiles of representative genes across the 10 clusters, including hepatic stellate cells, T cells, macrophages, hepatocytes, cholangiocytes, prtal vein ECs, periportal LSECs, midzonal LSECs, pericentral LSECs, and central vein ECs (C) Representative cluster-specific markers for five major liver endothelial subpopulations. (D) Spatial schematic illustrating the anatomical position of the Periportal Lamellar Complex (PLC), located exclusively between the portal vein ECs and periportal LSECs. (E) Venn diagram showing overlapping top 20 highly expressed genes between prtal vein ECs and periportal LSECs. (F) Multiplex immunofluorescence images showing CD31 (yellow), Sca-1 (green), and CD34 (red) expression in distinct hepatic endothelial zones, including prtal vein ECs, midzonal LSECs, and central vein ECs. The PLC structure is demarcated by dashed white arrows. Scale bar: 40 μm. (G) Classification of CD34+Sca-1+ double-positive versus double-negative endothelial subpopulations. (H) Volcano plot illustrating differentially expressed hematopoietic-associated genes in CD34+Sca-1+ cells. The x-axis shows log2 fold-change (log2FC); the y-axis shows −log10 adjusted P value. Significantly upregulated genes are located in the upper right quadrant. (I) Gene Ontology enrichment analysis of genes upregulated in CD34+Sca-1+ ECs, highlighting functional categories such as stem cell development, differentiation, and stem cell maintenance. (J) The schematic illustration depicts the spatial localization of Sca-1, a marker for mesenchymal or hematopoietic stem cells, and CD34+Sca-1+CD31+ endothelial cells as the trunk of the Periportal Lamellar Complex (PLC).

Earlier version: The PLC structures are anatomically located within the transitional zone between portal vein ECs and periportal LSECs(Figure 4D). Most of the characteristic genes within this region exhibited a continuous gradient expression pattern, rather than displaying a distinct binary classification between portal vein ECs and Periportal LSECs. Based on this spatial continuum, we hypothesized that PLC endothelial cells might display a hybrid transcriptional signature, co-expressing marker genes derived from both adjacent endothelial populations (Figures 4C and 4D).To test this, we performed Venn diagram intersection analysis combined with differential gene expression screening, identifying a set of candidate genes potentially specific to PLC endothelial cells (Figure 4E). Subsequent validation by CmTSA staining confirmed that CD31, a pan-endothelial marker, reliably delineated the spatial boundaries of portal vein ECs, PLC endothelium, and periportal LSECs (Figure 4F). Notably, among the intersection genes, CD34 and Sca-1 (Ly6a) were co-expressed exclusively within the PLC structures surrounding the portal vein, but absent from both central vein ECs and midzonal LSECs (Figure 4F).Additionally, we examined the spatial distribution of other endothelial-associated markers in two-dimensional liver sections. CD36 was found to specifically label

Revised version: The PLC structures are anatomically located within the transitional zone between portal vein ECs and periportal LSECs (Figure 4D). Most characteristic genes within this region exhibited a continuous gradient expression pattern rather than a distinct binary classification. Based on this spatial continuum, we hypothesized that PLC endothelial cells might display a hybrid transcriptional signature, co-expressing marker genes from both adjacent populations (Figure 4C, D). Venn diagram intersection analysis combined with differential gene expression screening identified a set of candidate genes potentially specific to PLC endothelial cells (Figure 4E). TSA staining confirmed that CD31, a pan-endothelial marker, reliably delineated the spatial boundaries of portal vein ECs, PLC endothelium, and periportal LSECs (Figure 4F). Notably, CD34 and Sca-1 (Ly6a) were co-expressed exclusively within PLC structures surrounding the portal vein, but absent from central vein ECs and midzonal LSECs (Figure 4F).

Added in the revised version: We further examined the spatial distribution of other endothelial-associated markers in two-dimensional liver sections. CD36 specifically labeled periportal LSECs, while PDGFRβ was broadly expressed across multiple intrahepatic endothelial subtypes (Figure 4—figure supplement 1C). Comparative analysis indicated that CD34+Sca-1+ expression serves as a superior marker set for identifying PLC endothelial cells. Based on this, we sorted CD34+Sca-1+ endothelial populations from the total liver EC pool (Figure 4G). Projection mapping analysis showed that CD34 and Sca-1 were primarily distributed within portal vein ECs, periportal LSECs, and midzonal LSECs (Figure 4—figure supplement 1D). Given that CD34 and Sca-1 are classical markers of mesenchymal stem cells and intrahepatic hematopoietic stem cells (HSCs), these observations suggest potential stem/progenitor-like properties of PLC-associated endothelial cells (Agrawal, Mehatre, & Khurana, 2024; Morcos et al., 2017; B. E. Petersen et al., 2003; Suskind & Muench, 2004; Taniguchi, Toyoshima, Fukao, & Nakauchi, 1996).

Earlier version: Periportal LSECs, while PDGFRβ was broadly expressed across multiple intrahepatic endothelial subtypes(Figure S4D). Comparative analysis indicated that the combination of CD34⁺Sca-1⁺ expression served as a superior marker set for identifying PLC endothelial cells.Based on this, we sorted CD34⁺Sca-1⁺ endothelial populations from the total liver endothelial cell pool (Figure 4G). Projection mapping analysis demonstrated that CD34 and Sca-1 were primarily distributed within Portal vein ECs, Periportal LSECs, and Midzonal LSECs (Figure S4C). Notably, CD34 and Sca-1 are also classical markers of mesenchymal stem cells and intrahepatic hematopoietic stem cells (HSCs), indicating potential stem/progenitor-like properties of PLC-associated endothelial cells[25–29]. Further analysis revealed that the CD34⁺Sca-1⁺ endothelial cell population, which marks the PLC structures, exhibited significantly upregulated expression of hematopoietic stem cell (HSC)-associated genes. Notably, Mecom — a critical transcriptional regulator essential for HSC self-renewal and highly expressed in human and approximately 60% of murine long-term HSCs[30, 31]— was prominently enriched within this population. In addition, PDGFRα, which together with Sca-1 serves as a classical marker for mesenchymal stem cell (MSC) isolation and may label important mesenchymal progenitor cell populations known to promote hepatocyte differentiation via direct contact and growth factor secretion[32, 33], was also highly expressed (Figure 4H).Gene Ontology (GO) enrichment analysis further demonstrated that differentially expressed genes in the CD34⁺Sca-1⁺ population were significantly enriched in pathways associated with HSC proliferation, lineage differentiation, mesenchymal cell development, and hepatic cell fate determination (Figure 4I). These findings suggest that this unique endothelial subset possesses dual regulatory functions in both metabolic and hematopoietic modulation within the periportal region. Finally, three-dimensional structural reconstruction (Figure 4J) revealed that the main trunk of the PLC is composed of CD34⁺Sca-1⁺CD31⁺ endothelial cells. In addition, a distinct “hematopoietic stem cell–vascular” niche unit formed by CD34⁺Sca-1⁺ double- positive cells was identified at the basal region of the PLC structure. This finding suggests that PLC endothelial cells not only regulate the periportal microcirculatory blood flow, but also establish a specialized microenvironment that supports periportal hematopoietic regulation, contributing to stem cell recruitment, vascular homeostasis, and tissue repair.

Revised version: Further analysis revealed that the CD34+Sca-1+ endothelial population marking the PLC structures exhibited significantly upregulated expression of HSC-associated genes (Figure 4H). Mecom, a critical transcriptional regulator essential for HSC self-renewal and highly expressed in human and approximately 60% of murine long-term HSCs (Christodoulou et al., 2020; Y. Zhang et al., 2011), was prominently enriched within this population. Additionally, PDGFRα, a classical marker for mesenchymal stem cell (MSC) isolation and a potential label for mesenchymal progenitor cells that promote hepatocyte differentiation via direct contact and growth factor secretion (Kikuchi & Monga, 2015; Owen, Patten, Vigneswara, Frampton, & Newsome, 2022), was also highly expressed (Figure 4H). Gene Ontology (GO) enrichment analysis further revealed that differentially expressed genes in the CD34+Sca-1+ population were significantly enriched in pathways associated with the development, proliferation, and differentiation of multiple stem cell types (including hematopoietic and mesenchymal stem cells), as well as in the maintenance of hematopoietic stem cell homeostasis (Figure 4I). These findings suggest that this unique endothelial cell subset in the periportal region may possess dual regulatory functions in both metabolic and hematopoietic modulation.

Added in the revised version: Finally, the main trunk of the PLC is composed of CD34+Sca-1+CD31+ endothelial cells (Figure 4J). A distinct “hematopoietic stem cell–vascular” niche was formed by CD34+Sca-1+ double-positive cells. This niche was located at the basal region of the PLC structure. These results suggest that PLC endothelial cells may not only regulate periportal microcirculatory blood flow but also help establish a specialized microenvironment that potentially supports periportal hematopoietic regulation, contributing to stem cell recruitment, vascular homeostasis, and tissue repair.

Earlier version: CD34⁺Sca-1⁺ Endothelium in the Periportal Lamellar Complex is potential Regulator of Spatial Patterning of Intrahepatic Bile Duct Branching during cirrhosis

Revised version: CD34+Sca-1+ endothelium in the periportal lamellar complex regulates spatial patterning of intrahepatic bile duct branching during liver fibrosis

Earlier version: Our 3D imaging showed that terminal bile duct branches are spatially colocalized with PLC structures around the main portal vein axis(Figures S5A-C). Volcano plot analysis identified multiple genes in the PLC endothelium associated with bile duct development and functional regulation, including JCAD, DLL4, and HES1 (Figure 5A).

Revised version: Our 3D imaging showed that terminal bile duct branches are spatially colocalized with PLC structures around the main portal vein axis (Figure 5—figure supplement 1A-B). Volcano plot re-analysis identified multiple genes in the PLC endothelium associated with bile duct development and functional regulation, including JCAD, DLL4, and HES1 (Figure 5A). Gene Ontology enrichment analysis confirmed that differentially expressed genes in the CD34+Sca-1+ endothelium of PLC are enriched in pathways regulating epithelial morphogenesis and branching morphogenesis of epithelial tubes, suggesting that the PLC may play multifunctional roles in bile duct branching and hepatic microenvironmental homeostasis (Figure 5B).

Earlier version: CD34⁺Sca-1⁺ Endothelium in the Periportal Lamellar Complex is potential Regulator of Spatial Patterning of Intrahepatic Bile Duct Branching during cirrhosis.

Revised version: CD34+Sca-1+ endothelium in the periportal lamellar complex regulates spatial patterning of intrahepatic bile duct branching during liver fibrosis.

Earlier version: (A) Volcano plot of differentially expressed bile duct-related genes in CD34⁺Sca-1⁺ double-positive cells. The x-axis represents log2 fold change (log2FC), and the y-axis represents -log10 adjusted P value (-log10(P-adjust)). Genes significantly upregulated are located in the upper right quadrant. (B) Functional enrichment analysis of upregulated genes in CD34⁺Sca-1⁺ cells, showing enriched categories such as epithelial morphogenesis and branching morphogenesis of epithelial tubes, represented by -log10(P value). (C) Visualization of portal vein labeled with MCNP-Pink and bile duct epithelial cells stained for CK19 (brown) by 3D DAB immunohistochemistry in control and CCl₄ 6-week fibrotic mice. Scale bar: 200 μm. (D) Quantification of the distance that bile duct termini extend from portal vein surfaces into hepatic parenchyma in control and CCl₄ 6-week mice, presented as mean ± SD (control n=20, CCl₄ 6-week n=18). (E) Multiplex immunofluorescence showing expression and spatial distribution of CD31 (yellow), Sca-1 (green), and CK19 (red) in control and fibrotic models (CCl₄ 3-week and 6-week). White arrows in CK19 single-channel magnified images indicate bile duct termini. Scale bar: 50 μm. (F) Quantitative measurement of bile duct termini extension distances along PLC structures into hepatic parenchyma in control, early (CCl₄ 3-week), and late (CCl₄ 6-week) fibrosis mice. Data represent mean ± SD (n=5 per group). Statistical significance determined by one-way ANOVA with Tukey’s multiple comparisons test; *P < 0.05, **P < 0.01, ****P < 0.0001. (G) Volcano plot of differentially expressed bile duct-related genes in CD34⁺Sca-1⁺ cells from fibrotic livers compared to controls. Axes as in (A). Significantly upregulated genes are in the upper right quadrant. (H) Schematic illustration showing spatial localization of CK19 and CD31 within the PLC structure.

Revised version: (A) Volcano plot of differentially expressed bile duct-related genes in CD34+Sca-1+ double-positive cells. The x-axis represents log2 fold change (log2FC), and the y-axis represents −log10 adjusted P value (−log10(P-adjust)). Genes significantly upregulated are located in the upper right quadrant. (B) Pathway enrichment analysis of upregulated genes in CD34+Sca-1+ cells showed enrichment in categories such as epithelial morphogenesis and branching morphogenesis of epithelial tubes, represented as −log10(P value). (C) Visualization of portal vein labeled with MCNP-Pink and bile duct epithelial cells stained for CK19 (brown) by 3D DAB immunohistochemistry in control and CCl4 6-week fibrotic mice. Scale bar: 200 μm. (D) Quantification of the distance that bile duct termini extend from portal vein surfaces into hepatic parenchyma in control and CCl4 6-week mice, presented as mean ± SD (control n=20, CCl4 6-week n=18). (E) Multiplex immunofluorescence showing expression and spatial distribution of CD31 (yellow), Sca-1 (green), and CK19 (red) in control and fibrotic models (CCl4 3-week and 6-week). White arrows in CK19 single-channel magnified images indicate bile duct termini. Scale bar: 50 μm. (F) Quantitative measurement of bile duct termini extension distances along PLC structures into hepatic parenchyma in control, early (CCl4 3-week), and late (CCl4 6-week) fibrosis mice. Data represent mean ± SD (n=5 per group). Statistical significance determined by one-way ANOVA with Tukey’s multiple comparisons test; *P < 0.05, **P < 0.01, ****P < 0.0001. (G) Volcano plot of differentially expressed bile duct-related genes in CD34+Sca-1+ cells from fibrotic livers compared to controls. Axes as in (A). Significantly upregulated genes are in the upper right quadrant. (H) Schematic illustration showing spatial localization of CK19 and CD31 within the PLC structure.

Removed in the revised version: Gene Ontology enrichment analysis confirmed that differentially expressed genes in the CD34⁺Sca-1⁺ endothelium of PLC are enriched in pathways regulating epithelial morphogenesis and branching morphogenesis of epithelial tubes, suggesting multifunctional roles of PLC in bile duct branching and hepatic microenvironment homeostasis (Figure 5B)

Earlier version: During liver fibrosis progression, proliferation and infiltration of terminal bile duct branches (Canals of Hering, CoH) into the hepatic parenchyma represent a core feature of the intrahepatic ductular reaction (DR), critically influencing disease progression, liver regeneration, and carcinogenesis[34–36]. To investigate the potential regulatory role of the PLC structures in bile duct growth, we established a CCl₄-induced mouse model of liver fibrosis/cirrhosis. The results showed:(1)Using MCNP-Pink to label the portal vein combined with three-dimensional DAB immunohistochemical staining for CK19 to visualize bile duct epithelial cells, extended depth-of-field brightfield imaging revealed that in control mice, bile duct termini localized adjacent to the portal vein, whereas after 3 or 6 weeks of CCl₄ treatment, bile duct termini significantly extended into the hepatic parenchyma by 100–300 μm (Figures 5C and 5D), with a notable increase in bile duct branch area (Figure S5C). (2)Multiplex CmTSA 3D imaging demonstrated that under fibrotic conditions, PLC structures became elongated and extended toward the lobular parenchyma. Concurrently, the CD34⁺Sca-1⁺ endothelial cell population of PLC, showed a significant increase during liver fibrosis progression (Figure S5D). (3)Further 3D staining confirmed that under fibrotic conditions, bile ducts progressively grew and branching along the expanded PLC structures and infiltrated deeper into the hepatic lobules, forming extensive terminal bile duct branches within the parenchyma. (4)With advancing fibrosis, the extension of CK19⁺ bile duct termini into the parenchyma was markedly greater after 6 weeks of CCl₄ treatment compared to 3 weeks, suggesting a coordinated expansion of bile duct branches alongside PLC structures that correlates with fibrosis severity (Figures 5E and 5F).

Revised version: During liver fibrosis progression, proliferation and infiltration of terminal bile duct branches (Canals of Hering) into the hepatic parenchyma represent a core feature of the intrahepatic ductular reaction (DR), critically influencing disease progression, liver regeneration, and carcinogenesis (Carpino et al., 2018; Clouston et al., 2005; Strazzabosco & Fabris, 2012). To investigate the potential regulatory role of the PLC structures in bile duct growth, we established a CCl4-induced mouse model of liver fibrosis. The results showed: (1)using MCNP-Pink to label the portal vein combined with three-dimensional DAB immunohistochemical staining for CK19 to visualize bile duct epithelial cells, EDOF imaging revealed that in control mice, bile duct termini localized adjacent to the portal vein, whereas after 3 or 6 weeks of CCl4 treatment, bile duct termini significantly extended into the hepatic parenchyma by 100–300 μm (Figure 5C, D), with a notable increase in bile duct branch area (Figure 5—figure supplement 1C). (2)multiplex TSA 3D imaging demonstrated that under fibrotic conditions, PLC structures became elongated and extended toward the lobular parenchyma. Concurrently, the CD34+Sca-1+ endothelial cell population of PLC showed a significant increase during liver fibrosis progression (Figure 5—figure supplement 1D). (3)further 3D staining confirmed that under fibrotic conditions, bile ducts progressively grew and branching along the expanded PLC structures and infiltrated deeper into the hepatic lobules, forming extensive terminal bile duct branches within the parenchyma. (4)with advancing fibrosis, the extension of CK19+ bile duct termini into the parenchyma was markedly greater after 6 weeks of CCl4 treatment compared to 3 weeks, suggesting a coordinated expansion of bile duct branches alongside PLC structures that correlates with fibrosis severity (Figure 5E, F).

Earlier version: Moreover, single-cell transcriptomic analysis revealed significant upregulation of bile duct-related genes in the CD34⁺Sca-1⁺ endothelium of PLC in cirrhotic liver, with notably high expression of Lgals1 (Galectin-1) and HGF(Figure 5G). Previous studies have shown that Galectin-1 is absent in normal liver parenchyma but highly expressed in intrahepatic cholangiocarcinoma (ICC), correlating with tumor dedifferentiation and invasion[37, 38]. Additionally, hepatocyte growth factor (HGF), particularly in combination with epidermal growth factor (EGF) in 3D cultures, promotes hepatic progenitor cells to form bile duct-polarized cystic structures[39]. Together, these findings suggest the PLC endothelium may act as a key regulator of bile duct branching and fibrotic microenvironment remodeling in liver cirrhosis.

Revised version: Moreover, single-cell transcriptomic re-analysis revealed significant upregulation of bile duct-related genes in the CD34+Sca-1+ endothelium of PLC in fibrotic liver, with notably high expression of Lgals1 (Galectin-1) and Hgf (Figure 5G). Previous studies have shown that Galectin-1 is absent in normal liver parenchyma but highly expressed in intrahepatic cholangiocarcinoma (ICC), correlating with tumor dedifferentiation and invasion (Bacigalupo, Manzi, Rabinovich, & Troncoso, 2013; Shimonishi et al., 2001). Additionally, hepatocyte growth factor (HGF), particularly in combination with epidermal growth factor (EGF) in 3D cultures, promotes hepatic progenitor cells to form bile duct-polarized cystic structures (N. Tanimizu, Miyajima, & Mostov, 2007). Together, these findings suggest the PLC endothelium may act as a key regulator of bile duct branching and fibrotic microenvironment remodeling in liver fibrosis.

Earlier version: Collectively, our results demonstrate that the PLC, situated between the portal vein and periportal sinusoidal endothelium, constitutes a critical vascular microenvironmental unit. It not only guides bile duct branching during development but also, through its basal CD34⁺Sca-1⁺ double-positive endothelial cells, orchestrates bile duct epithelial proliferation, branching morphogenesis, and bile acid transport homeostasis via multiple signaling pathways. Particularly during liver fibrosis progression, the PLC exhibits dynamic structural extension correlating with fibrosis severity, serving as a spatial scaffold facilitating terminal bile duct migration and expansion into the hepatic parenchyma (Figure 5H). These findings highlight the PLC endothelial cell population and the vascular-bile duct interface as key regulatory hubs in bile duct regeneration, tissue repair, and pathological remodeling, providing novel cellular and molecular insights for understanding bile duct-related diseases such as ductular reaction, cholangiocarcinoma, and cholestatic disorders, and offering potential targets for therapeutic intervention.

Revised version: Collectively, our results demonstrate that the PLC, situated between the portal vein and periportal sinusoidal endothelium, constitutes a critical vascular microenvironmental unit. It may not only colocalize with bile duct branches under normal physiological conditions, but also through its basal CD34+Sca-1+ double-positive endothelial cells, potentially orchestrate bile duct epithelial proliferation, branching morphogenesis, and bile acid transport homeostasis via multiple signaling pathways. Particularly during liver fibrosis progression, the PLC exhibits dynamic structural extension, serving as a spatial scaffold facilitating terminal bile duct migration and expansion into the hepatic parenchyma (Figure 5H). These findings highlight the PLC endothelial cell population and the vascular-bile duct interface as key regulatory hubs in bile duct regeneration, tissue repair, and pathological remodeling, providing novel cellular and molecular insights for understanding bile duct-related diseasessuch as ductular reaction, cholangiocarcinoma, and cholestatic disorders, and offering potential targets for therapeutic intervention.

Earlier version: CD34⁺Sca-1⁺ Endothelium in the Periportal Lamellar Complex is a Potential Neurovascular Niche Regulating Hepatic Autonomic Nerve Patterning in Cirrhosis

Revised version: CD34+Sca-1+ endothelium in the periportal lamellar complex forms a neurovascular niche regulating hepatic autonomic nerve patterning

Added in the revised version: In this study, we examined the spatial distribution of hepatic nerves in mice using immunofluorescence staining and found that nerve fibers were almost entirely confined to the portal vein region (Figure 6—figure supplement 1A). Notably, this distribution pattern differs markedly from that observed in humans—previous studies have shown that, in addition to the portal area, hepatic nerves in humans are also present around the central vein, within the interlobular septa, and in the connective tissue of the parenchyma (Miller et al., 2021; Yi, la Fleur, Fliers, & Kalsbeek, 2010).

Removed in the revised version: The liver receives both afferent and efferent innervation from sympathetic and parasympathetic nerve fibers. Although hepatic nerve distribution exhibits species- specific differences, In the human liver, nerve fibers are predominantly localized around the portal vein and central vein regions, and terminal nerve cells can also be observed within the interlobular septa and the connective tissue of the parenchyma. In contrast, in mice and rats, hepatic innervation is almost exclusively confined to the portal vein area (Figure S6A)[6, 40]. How the autonomic nervous system regulates liver function in mice despite the apparent absence of substantive nerve fiber invasion into the parenchyma remains unclear.

Removed in the revised version: In this study, we identified that the CD34⁺Sca-1⁺ endothelium of PLC exhibits prominent neurodevelopmental molecular features. Differential gene expression analysis revealed significant upregulation of multiple genes associated with neurogenesis and axon guidance within this population (Figure 6A), including Nrg1—a key regulator of Schwann cell proliferation, myelination, and sheath formation—and ADGRG6 (Gpr126), an essential factor for myelination in the mammalian peripheral nervous system, required for Schwann cell differentiation and peripheral nerve development in mice[41]. Subsequent GO enrichment analysis of differentially expressed genes (DEGs) in this cell cluster demonstrated significant activation of pathways related to neuronal projection and central nervous system differentiation (Figure 6B).

Earlier version: CD34⁺Sca-1⁺ Endothelium in the Periportal Lamellar Complex is a Potential Neurovascular Niche Regulating Hepatic Autonomic Nerve Patterning in Cirrhosis

Revised version: CD34+Sca-1+ endothelium in the periportal lamellar complex forms a neurovascular niche regulating hepatic autonomic nerve patterning.

Earlier version: (A) Differential expression analysis of nerve-related genes in CD34⁺Sca-1⁺ endothelial cells. The x-axis indicates log2 fold change (log2FC), and the y-axis represents –log10 adjusted P value (–log10(P-adjust)). Significantly upregulated genes are located in the upper right quadrant. (B) Gene ontology (GO) enrichment analysis of upregulated genes in CD34⁺Sca-1⁺ cells, showing enrichment in functional categories such as semaphorin-plexin mediated axon guidance, regulation of neuronal projection regeneration, and modulation of postsynaptic neurotransmitter receptor endocytosis. Enrichment significance is indicated by –log10(P-value). (C) Multiplex immunofluorescence staining showing tyrosine hydroxylase (TH, green) labeling sympathetic nerves and CD31 (yellow) labeling portal vein endothelium. Scale bar: 20 μm. (D) Distribution of CD31 (yellow) and TH (green) expression in control and CCl₄-induced liver fibrosis models at week 3 and week 6, visualized by multiplex immunofluorescence. In the green channel images, the white arrows marks the terminal location of TH-positive sympathetic nerve endings. Scale bar: 50 μm. (E) Quantification of the distance from sympathetic nerve endings to the hepatic parenchyma along PLC structures in control, early fibrosis (CCl₄-3 weeks), and advanced fibrosis (CCl₄-6 weeks) mice. Data are presented as mean ± standard deviation (Mean ± SD), n=5 per group. Statistical analysis was performed using one-way ANOVA followed by Tukey’s multiple comparison test; *P < 0.05, **P < 0.01, ****P < 0.0001. (F) Differential expression analysis of nerve-related genes between CD34⁺Sca-1⁺ cells from cirrhotic and control livers. The x-axis indicates log2 fold change (log2FC), and the y-axis represents –log10 adjusted P value (–log10(P- adjust)). Significantly upregulated genes are located in the upper right quadrant. (G) Schematic diagram illustrating the spatial localization of TH-positive sympathetic nerves and CD31-positive endothelial cells within PLC structures.

Revised version: (A) Differential expression analysis of nerve-related genes in CD34+Sca-1+ endothelial cells. The x-axis indicates log2 fold change (log2FC), and the y-axis represents –log10 adjusted P value (–log10(P-adjust)). Significantly upregulated genes are located in the upper right quadrant. (B) Gene ontology (GO) enrichment analysis of upregulated genes in CD34+Sca-1+ cells, showing enrichment in functional categories such as semaphorin-plexin mediated axon guidance, regulation of neuronal projection regeneration, and modulation of postsynaptic neurotransmitter receptor endocytosis. Enrichment significance is indicated by –log10(P-value). (C) Multiplex immunofluorescence staining showing tyrosine hydroxylase (TH, green) labeling sympathetic nerves and CD31 (yellow) labeling portal vein endothelium. Scale bar: 20 μm. (D) Distribution of CD31 (yellow) and TH (green) expression in control and CCl4-induced liver fibrosis models at week 3 and week 6, visualized by multiplex immunofluorescence. In the green channel images, the white arrows marks the terminal location of TH-positive sympathetic nerve endings. Scale bar: 50 μm. (E) Quantification of the distance from sympathetic nerve endings to the hepatic parenchyma along PLC structures in control, early fibrosis (CCl4-3 weeks), and advanced fibrosis (CCl4-6 weeks) mice. Data are presented as mean ± standard deviation (Mean ± SD), n=5 per group. Statistical analysis was performed using one-way ANOVA followed by Tukey’s multiple comparison test; *P < 0.05, **P < 0.01, ****P < 0.0001. (F) Differential expression analysis of nerve-related genes between CD34+Sca-1+ cells from fibrosis and control livers. The x-axis indicates log2 fold change (log2FC), and the y-axis represents –log10 adjusted P value (–log10(P-adjust)). Significantly upregulated genes are located in the upper right quadrant. (G) Schematic diagram illustrating the spatial localization of TH-positive sympathetic nerves and CD31-positive endothelial cells within PLC structures.

Added in the revised version: Previous studies have further explained the physiological basis for this difference: even among species that differ in parenchymal sympathetic innervation (i.e., species with or without direct sympathetic input), their sympathetic efferent regulatory functions may remain comparable (Beckh, Fuchs, Ballé, & Jungermann, 1990). This is because signals released from aminergic and peptidergic nerve terminals can be transmitted to hepatocytes in the form of electrical impulses through intercellular gap junctions (Hertzberg & Gilula, 1979; Jensen, Alpini, & Glaser, 2013; Seseke, Gardemann, & Jungermann, 1992; Taher, Farr, & Adeli, 2017). However, the scarcity of nerve fibers within the mouse hepatic parenchyma implies that the mechanisms of autonomic regulation of liver function in mice may differ from those in humans, prompting us to further investigate the potential role of PLC endothelial cells in this process.

Added in the revised version: We further identified that the CD34+Sca-1+ endothelium of the PLC exhibits prominent neurodevelopmental molecular features. Differential gene expression analysis revealed significant upregulation of multiple genes associated with neurogenesis and axon guidance within this population (Figure 6A), including Nrg1. Previous studies have shown that NRG-1 is present in intrahepatic nerves and colocalizes with neuronal nitric oxide synthase (nNOS) and vasoactive intestinal peptide (VIP). Exposure to high doses of bisphenol A (BPA) significantly increases the density of NRG-1–immunoreactive fibers in intrahepatic nerves and enhances the colocalization of NRG-1 with VIP, suggesting that NRG-1 in intrahepatic nerves may play a role in BPA-induced liver injury and/or liver regeneration(Szymańska, Makowska, Całka, & Gonkowski, 2020); ADGRG6 (GPR126) might regulate liver innervation through non-myelination-related mechanisms. Indeed, GPR126 has been shown to mediate Schwann cell–dependent and independent processes during peripheral nerve injury and repair, influencing axon guidance, mechanosensing, and ECM remodeling (Mogha et al., 2016; Monk et al., 2009; S. C. Petersen et al., 2015). These findings suggest that GPR126 may participate in neural–microenvironment communication within the periportal region (Monk, Oshima, Jörs, Heller, & Talbot, 2011). Subsequent GO enrichment analysis of differentially expressed genes (DEGs) in this cell cluster demonstrated significant activation of pathways related to neuronal projection and central nervous system differentiation (Figure 6B).

Earlier version: Previous studies have confirmed that the highest density of hepatic parenchymal innervation originates from postganglionic sympathetic neurons located in the celiac ganglion[42]. Tyrosine hydroxylase (TH) immunohistochemistry is the preferred method to visualize noradrenergic sympathetic innervation of hepatocytes and vasculature in both rodents and humans[43]. In this study, confocal three-dimensional imaging revealed that TH⁺ sympathetic nerve fibers form a characteristic “neuronal bead-chain” network along the PLC structure, with nerve terminals terminating at the PLC base and around vascular bundles, suggesting a potential regulatory role of the PLC in controlling functional zones within the hepatic lobule (Figures 6C, S6B and S6C).

Revised version: Previous studies have confirmed that the highest density of hepatic parenchymal innervation originates from postganglionic sympathetic neurons located in the celiac ganglion(Mizuno & Ueno, 2017). Tyrosine hydroxylase (TH) immunohistochemistry is the preferred method to visualize noradrenergic sympathetic innervation of hepatocytes and vasculature in both rodents and humans(Adori et al., 2021). In this study, confocal three-dimensional imaging revealed that TH+ sympathetic nerve fibers form a characteristic “neuronal bead-chain” network along the PLC structure, with nerve terminals terminating at the PLC base and around vascular bundles, suggesting a potential regulatory role of the PLC in controlling functional zones within the hepatic lobule (Figure 6C, Figure 6—figure supplement 1B, C).

Earlier version: Further validation using a CCl₄-induced mouse liver fibrosis model demonstrated that: (1) during fibrosis progression, sympathetic nerve fibers progressively extend along the PLC structure, accompanied by nerve terminal branching and infiltration into the hepatic parenchyma; (2) after 6 weeks of CCl₄ treatment, TH⁺ nerve fibers extended significantly farther into the hepatic parenchyma compared to 3 weeks, indicating a coordinated expansion trend of bile duct branches, PLC structures, and sympathetic nerve networks closely associated with fibrosis severity (Figures 6D and 6E).

Revised version: Further validation using a CCl4-induced mouse liver fibrosis model demonstrated that: (1) during fibrosis progression, sympathetic nerve fibers progressively extend along the PLC structure, accompanied by nerve terminal branching and infiltration into the hepatic parenchyma; (2) after 6 weeks of CCl4 treatment, TH+ nerve fibers extended significantly farther into the hepatic parenchyma compared to 3 weeks, indicating a coordinated expansion trend of bile duct branches, PLC structures, and sympathetic nerve networks closely associated with fibrosis severity (Figure 6D, E).

Earlier version: Furthermore, single-cell transcriptomic volcano plot analysis revealed significant upregulation of neural-related genes within the CD34⁺Sca-1⁺ endothelium of PLC in the cirrhotic group, notably including EDNRB (Endothelin receptor type B), which plays a critical role in the development of neural crest-derived autonomic nervous system components, including the liver[39]. Nestin, a type VI intermediate filament protein predominantly expressed in the central nervous system, has recently been implicated in tissue homeostasis during wound healing. For example, in traumatic CNS injury, Nestin is induced in reactive astrocytes contributing to glial scar formation. Additionally,

Revised version: Furthermore, single-cell transcriptomic volcano plot re-analysis revealed significant upregulation of neural-related genes within the CD34+Sca-1+ endothelium of PLC in the fibrosis group, notably including EDNRB (Endothelin receptor type B), which plays a critical role in the development of neural crest-derived autonomic nervous system components, including the liver(N. Tanimizu et al., 2007). Nestin, a type VI intermediate filament protein predominantly expressed in the central nervous system, has recently been implicated in tissue homeostasis during wound healing. For example, in traumatic CNS injury, Nestin is induced in reactive astrocytes contributing to glial scar formation. Additionally, Nestin expression correlates with the severity of renal tubulointerstitial fibrosis. In the healthy adult liver, Nestin is nearly absent but is upregulated following acute or chronic hepatic injury, indicating its potential role in liver regeneration(Figure 6F) (Ahmed, Kunisaki, Merad, & Frenette, 2013).

Added in the revised version: To further validate the reliability of our findings, we incorporated another high-depth single-cell RNA sequencing (scRNA-seq) dataset for cross-analysis and re-exploration(Pietilä et al., 2025). This dataset provides a detailed molecular atlas of the major cell types in the adult mouse liver. In their original analysis, Pietilä et al. annotated cluster #e2 as portal vein endothelial cells based on well-established cell-type-specific markers, such as Ly6a/Sca-1 and Cd34 for portal vein endothelium. This annotation closely aligns with the CD34+Sca-1+ endothelial cell population we identified in the PLC region, providing cross-dataset preliminary validation for the existence and core molecular features of this endothelial subpopulation.

Added in the revised version: Building on this, we performed targeted re-analysis of the reference dataset: by selecting the CD34+Sca-1+ double-positive cell population, we found that their highly expressed gene profile closely matched our own database results, notably including genes related to hematopoietic stem cells (e.g., ACE, Mecom), bile duct-associated genes (e.g., DLL4, JCAD), and neural-associated genes (e.g., NRG1,Adgrg6), all of which were significantly upregulated (Figure 6—figure supplement 1D—source data 1).

Added in the revised version: Meanwhile, gene ontology (GO) enrichment analysis of this cell population further confirmed that pathways related to hematopoietic stem cell function regulation (including stem cell development and differentiation), neural development (e.g., neural development), and bile duct-related processes (epithelial tube formation pathway) were all significantly upregulated (Figure 6—figure supplement 1E). This result echoes the gene expression characteristics, as it not only strengthens the functional associations between CD34+Sca-1+ endothelial cells in the PLC region and the hematopoietic, bile duct, and nervous systems, but also provides molecular evidence for their multipotential regulatory properties

Earlier version: Nestin expression correlates with the severity of renal tubulointerstitial fibrosis. In the healthy adult liver, Nestin is nearly absent but is upregulated following acute or chronic hepatic injury, indicating its potential role in liver regeneration(Figure 6F)[44]. These findings suggest that this specific endothelial subpopulation of PLC may participate in neural branching development and fibrotic microenvironment remodeling, exhibiting multifunctional regulatory potential(Figure 6G).

Revised version: These findings suggest that this specific endothelial subpopulation of PLC may participate in neural branching development and fibrotic microenvironment remodeling, exhibiting multifunctional regulatory potential(Figure 6G).

Discussion

Added in the revised version: Accurate three-dimensional mapping of the liver is crucial for understanding its function and disease mechanisms. Conventional 2D histology and flow cytometry lose spatial information, particularly in the liver’s complex ductal-vascular systems. To address technical challenges in 3D liver imaging—namely, tissue opacity due to heme, lipids, and dense connective tissue, and trade-offs between resolution and field of view—we developed the Liver-CUBIC system. By optimizing H2O2 bleaching and urea concentration, tissue transparency was significantly improved, while multicolor metal compound nanoparticles (MCNPs) enabled four-channel labeling of portal veins, hepatic arteries, central veins, and bile ducts. This approach allows high-resolution 3D visualization of fine microstructures, including terminal bile duct arborization, portal vein branching, and central vein fenestrations, providing a robust platform for studying liver microcirculation (Hankeova et al., 2021; Peeters et al., 2018; Qi et al., 2019; Renier et al., 2014; Tainaka et al., 2014).

Removed in the revised version: Accurate three-dimensional mapping of organs in both physiological and pathological states is essential for understanding their function and disease mechanisms. Compared to flow cytometry and conventional two-dimensional histology, building a comprehensive multi-channel 3D atlas of the liver helps compensate for spatial information loss—especially important for its complex and heterogeneous ductal-vascular systems. This study addresses two major technical bottlenecks in 3D imaging of the mouse liver. First, although various tissue-clearing techniques (e.g., SeeDB, iDISCO, CUBIC, Clarity, and PEGASOS) have been applied to whole-organ imaging[16, 45, 46], he liver’s high content of heme, lipids, and dense connective tissue renders it significantly less transparent than organs like the brain.

Removed in the revised version: Second, current imaging platforms often face trade-offs between resolution and field of view. For example, CT and PCCT are limited to single-system visualization, while DUCT enables dual-system reconstruction but cannot resolve fine structures such as bile canaliculi due to its 5 μm/pixel resolution[9, 11].

Removed in the revised version: To overcome these challenges, we developed the Liver-CUBIC 3D visualization system. By optimizing H2O2-based bleaching and urea concentration, we significantly improved hepatic tissue transparency (total clearing time: 3.25 days; light transmittance increased by 20.12%). In parallel, we introduced multicolor metal compound nanoparticles (MCNPs) with excellent photostability and H2O2 resistance to achieve four-channel segmentation and labeling of the portal vein, hepatic artery, central vein, and bile ducts. This system enables, for the first time, 3D visualization of microstructures such as terminal bile duct arborization, fine portal vein branching patterns, and fenestrations in central venous endothelium—providing a high-resolution, multidimensional platform for liver microcirculation studies.

The labeling strategy relies on perfusing contrast-complementary MCNPs into multiple vascular and ductal compartments. By carefully controlling injection site, pressure, and volume—combined with 15 μL microinjection needles and post-perfusion vessel ligation—we ensured even dye distribution and structural fidelity. High-density image data can be simplified using dual-channel overlays (e.g., portal vein and bile duct), and multiple imaging modalities (e.g., stereomicroscopy, confocal, or extended-depth microscopy) can be chosen depending on the desired resolution or field of view. Moreover, the dyes demonstrate robust fluorescence under laser-scanning microscopy, enabling clear 3D structural imaging.

Added in the revised version: It should be noted that the current imaging thickness of our method is mostly around 200 μm, and it is not yet possible to achieve whole-liver tissue three-dimensional reconstruction at the level of light-sheet microscopy. However, through MCNP perfusion experiments on intact, unsectioned liver lobes, we have demonstrated dye distribution results covering the entire liver lobe (as shown in Figure 1—figure supplements1F): the images fully present the three-dimensional branching pattern and distribution range of the bile duct system within the liver lobe, clearly showing the continuous network structure from the main trunk to the terminal branches. This confirms that our study can also achieve overall visualization of the vascular system at the whole-organ level.

Added in the revised version: Although this system has significant advantages, it remains sensitive to anatomical variations of the liver, such as lobe overlap and heterogeneous vascular architecture. At vascular intersections, local perfusion may be uneven or excessively aggregated; therefore, injection strategies and perfusion parameters need to be adjusted according to liver size and vessel status to improve reproducibility and imaging quality. It should be noted that the results obtained with this method are primarily intended to visualize the overall and fine structures of the hepatic vascular and ductal systems, rather than to precisely reflect hemodynamic processes. In the future, we plan to further investigate the diffusion characteristics of the dyes within the liver microcirculation using in vivo perfusion or dynamic fluid models.

Removed in the revised version: Despite its strengths, the system remains sensitive to anatomical variation, such as liver lobe overlap and heterogeneous vascular architecture. Local perfusion may be inconsistent or excessively aggregated, especially at vascular intersections. Therefore, injection strategies and perfusion parameters must be adapted to liver size and vessel status to improve reproducibility and image quality.

Using this system, we identified a previously unrecognized structure located on the adventitia of the main portal vein, which we termed the Periportal Lamellar Complex (PLC). The PLC consists of multi-branched extensions radiating from the outer membrane of the main portal vein. Notably, it is absent in terminal vessels and other vascular systems, challenging the conventional “strict hierarchical bifurcation” model of vascular organization. The PLC is spatially associated with terminal branches of both bile ducts and hepatic arteries, suggesting possible roles in local microenvironmental regulation through physical or signaling interactions.

Earlier version: Multiplex immunostaining revealed that cholangiocytes extend along ZO-1⁺ bile canalicular networks but display an elongated, immature morphology without forming complete lumens (Fig. 3F). This observation aligns with previous findings that HNF4α⁺ hepatocytes and CK19⁺ cholangiocytes can jointly form bile-excreting conduits, supporting a guiding role for the PLC in bile duct positioning and morphogenesis[47].

Revised version: Multiplex immunostaining revealed that cholangiocytes extend along ZO-1+ bile canalicular networks but display an elongated, immature morphology without forming complete lumens (Fig. 3F). This observation aligns with previous findings that HNF4α+ hepatocytes and CK19+ cholangiocytes can jointly form bile-excreting conduits, supporting a guiding role for the PLC in bile duct positioning and morphogenesis(Naoki Tanimizu et al., 2021).

Earlier version: Morphologically, the PLC shares features with previously described telocytes (TCs)—a recently identified class of interstitial cells in the liver observed via transmission electron microscopy and immunohistochemistry as CD34⁺ cells with multiple long telopodes[48, 49]. However, functional and molecular characterization of TCs remains limited, and their direct relationship with the PLC is yet to be determined.

Revised version: Morphologically, the PLC shares features with previously described telocytes (TCs)—a recently identified class of interstitial cells in the liver observed via transmission electron microscopy and immunohistochemistry as CD34+ cells with multiple long telopodes(Xiao, Wang, Liu, & Yang, 2013; Xu et al., 2019). However, functional and molecular characterization of TCs remains limited, and their direct relationship with the PLC is yet to be determined.

Earlier version: Single-cell transcriptomic analysis revealed a CD34⁺Sca-1⁺ endothelial cell population enriched within the PLC that expressed hematopoiesis-related genes and stem cell regulatory pathways[22]. This suggests that the PLC may not only provide structural support but also serve as a perivascular stem cell niche specific to the portal region, potentially involved in hematopoiesis and tissue regeneration. Supporting this, Sca-1 expression is localized to the portal region during mouse liver development [50], and Nestin⁺NG2⁺ precursor cells in the fetal liver also concentrate near the portal vein, where they express key HSC-supportive factors such as SCF and CXCL12[44] — highlighting the PLC as a key component of the hepatic vascular-stem cell microenvironment.

Revised version: Single-cell transcriptomic analysis revealed a CD34+Sca-1+ endothelial cell population enriched within the PLC that expressed hematopoiesis-related genes and stem cell regulatory pathways(Su et al., 2021). This suggests that the PLC may not only provide structural support but also serve as a perivascular stem cell niche specific to the portal region, potentially involved in hematopoiesis and tissue regeneration. Supporting this, Sca-1 expression is localized to the portal region during mouse liver development (Gomez-Salinero et al., 2022), and Nestin+NG2+ precursor cells in the fetal liver also concentrate near the portal vein, where they express key HSC-supportive factors such as SCF and CXCL12(Ahmed et al., 2013) — highlighting the PLC as a key component of the hepatic vascular-stem cell microenvironment.

Earlier version: Further analysis showed significant upregulation of genes involved in neurodevelopment and axonal guidance in the CD34⁺Sca-1⁺ cluster, along with activation of neuronal signaling pathways. Immunostaining confirmed the presence of TH⁺ sympathetic nerve fibers wrapping around the PLC in a “beads-on-a-string” pattern (Fig. 6), consistent with a classic neurovascular unit[43]. Previous studies have shown that sympathetic nerves enter the liver along collagen fibers of Glisson’s capsule and interact with hepatic arteries, portal veins, and bile duct epithelium, supporting the PLC as a scaffold for intrahepatic neurovascular integration.

Revised version: Further analysis showed significant upregulation of genes involved in neurodevelopment and axonal guidance in the CD34+Sca-1+ cluster, along with activation of neuronal signaling pathways. Immunostaining confirmed the presence of TH+ sympathetic nerve fibers wrapping around the PLC in a “beads-on-a-string” pattern (Fig. 6), consistent with a classic neurovascular unit(Adori et al., 2021). Previous studies have shown that sympathetic nerves enter the liver along collagen fibers of Glisson’s capsule and interact with hepatic arteries, portal veins, and bile duct epithelium, supporting the PLC as a scaffold for intrahepatic neurovascular integration.

Earlier version: In CCl₄-induced hepatotoxic fibrosis models, injury and toxic metabolite accumulation are primarily restricted to the centrilobular region, while bile ducts extend linearly toward the central vein[51, 52]. We observed dynamic elongation of the PLC in parallel with fibrosis progression, indicating its role as a spatial scaffold for ductal expansion. This suggests potential involvement in ductular reactions and neural remodeling during cirrhosis.

Revised version: In CCl4-induced hepatotoxic fibrosis models, injury and toxic metabolite accumulation are primarily restricted to the centrilobular region, while bile ducts extend linearly toward the central vein(Kaneko, Kamimoto, Miyajima, & Itoh, 2015; Yang et al., 2021). We observed dynamic elongation of the PLC in parallel with fibrosis progression, indicating its role as a spatial scaffold for ductal expansion. This suggests potential involvement in ductular reactions and neural remodeling during fibrosis.

In conclusion, this study provides the first systematic characterization of the PLC as a structural and functional interface among hepatic vasculature, bile ducts, and nerves. The PLC serves as a directional scaffold for ductal growth, a specialized stem cell niche, and a potential site of neurovascular coupling. These findings provide novel insights into hepatic homeostasis, bile duct development, and fibrotic remodeling, and suggest that the PLC and its CD34+Sca-1+ endothelial population may represent promising therapeutic targets for diseases such as cholangiocarcinoma and primary sclerosing cholangitis.

Earlier version: STAR methods

Revised version: Methods

Removed in the revised version: Key resources table

Removed in the revised version: Experimental model and study participant details

Animals

Earlier version: All animal procedures were approved by the Institutional Animal Care and Use Committee of Sichuan University. Eight-week-old male wild-type mice (purchased from Beijing Huafukang Bioscience Co., Ltd.) were used to establish a progressive liver fibrosis model. Mice were intraperitoneally injected with carbon tetrachloride (CCl₄, 1 mL/kg body weight, CCl₄:olive oil = 1:4, v/v, Damas-beta, CAS:56-23-5) twice a week for 3 or 6 weeks (n=3 per group). Control mice received an equal volume of olive oil injection for 3 or 6 weeks (n=3 per group).

Revised version: All animal procedures were approved by the Institutional Animal Care and Use Committee of Sichuan University. Eight-week-old male wild-type mice (purchased from Beijing Huafukang Bioscience Co., Ltd.) were used to establish a progressive liver fibrosis model. Mice were intraperitoneally injected with carbon tetrachloride (CCl4, 1 mL/kg body weight, CCl4:olive oil = 1:4, v/v, Damas-beta, CAS:56-23-5) twice a week for 3 or 6 weeks (n=3 per group). Control mice received an equal volume of olive oil injection for 3 or 6 weeks (n=3 per group).

Added in the revised version: CCl4-Induced Liver Fibrosis Model

Added in the revised version: Step 1. Preparation of CCl4 Solution

Added in the revised version: Dissolve CCl4 in corn oil at a 1:3 (v/v) ratio under a fume hood and filter through a 0.2 μm filter into a 30 mL vial. Mice were injected intraperitoneally at 1 mL/kg to induce chronic liver injury.

Added in the revised version: Step 2. Mouse Weight and Injection Volume

Added in the revised version: Weigh each mouse to calculate the required injection volume.

Added in the revised version: Step 3. Loading Syringe

Added in the revised version: Draw the calculated CCl4 solution into a syringe and remove air bubbles.

Added in the revised version: Step 4. Mouse Positioning

Added in the revised version: Place the mouse on a textured surface, lower the head by ~30°, and identify the lower left or right abdominal quadrant.

Added in the revised version: Step 5. Injection

Added in the revised version: Insert the needle at a 45° angle (bevel up) approximately 4–5 mm into the abdomen and inject the calculated volume.

Added in the revised version: Step 6. Post-Injection Care

Added in the revised version: Return the mouse to its cage and monitor for any signs of distress. Dispose of syringes and needles safely; a new syringe and needle were used for each mouse.

Added in the revised version: Step 7. Injection Schedule

Added in the revised version: Administer injections twice weekly for 3 or 6 weeks (n = 3 per group). Control mice received equal volumes of corn oil.

Added in the revised version: Step 8. Tissue Collection

Added in the revised version: Forty-eight hours after the last injection, mice were sacrificed, livers were dissected, and samples were processed for histology or molecular analyses.

Added in the revised version: Human Liver Tissue Samples

Added in the revised version: The human liver tissue samples shown in Figure S6 were obtained from adjacent non-tumor liver tissues resected during surgical operations at West China Hospital, Sichuan University. All samples used were anonymized archived tissues, which were applied for scientific research in accordance with institutional ethical guidelines and did not involve any identifiable patient information. After being fixed in 10% neutral formalin for 24 hours, the tissues were routinely processed for paraffin embedding (FFPE), and sectioned into 4 μm-thick slices for immunostaining and fluorescence imaging.

Added in the revised version: Spectral Characteristics of Multicolor Metal Compound Nanoparticles (MCNPs)

Added in the revised version: The different types of metal compound nanoparticles (MCNPs) used in this study have distinct spectral characteristics, with specific parameters as follows:

Added in the revised version: Both MCNP-Green and MCNP-Yellow are AF488 spectrum-matched types, with an excitation wavelength of 495 nm and an emission wavelength of 519 nm, which are suitable for conventional fluorescence imaging channels.

Added in the revised version: MCNP-Pink is designed with a far-red spectrum, featuring an excitation wavelength of 561 nm and an emission wavelength of 640 nm. This design can effectively reduce the interference of tissue autofluorescence and improve the signal-to-noise ratio of deep-tissue imaging.

Added in the revised version: MCNP-Black has no fluorescent properties and only appears black under a bright-field microscope. It is mainly used to form a contrast with fluorescent MCNPs and assist in verifying the integrity of cavity perfusion.

Four-Channel Vascular Typing Labeling and Bile Duct High-Pressure Dye Perfusion in Mice

Step 1: Mice were euthanized by intraperitoneal (IP) injection of sodium pentobarbital at a dose of 100–150 mg/kg body weight.

Step 2: The chest muscles were incised to expose the heart. A perfusion needle was inserted into the left ventricle, ensuring entry into the cardiac cavity.

Step 3: Sequential perfusion was performed with 40 mL of pre-chilled PBS (5 mL/min) to clear blood, followed by 20 mL of 4% paraformaldehyde (5 mL/min) for in situ fixation.

Step 4: The abdominal skin and muscles were cut to the mid-axillary line to expose the liver. The intestines were gently displaced with sterile swabs to fully expose the common bile duct.

Step 5: According to experimental design, the hepatic artery, bile duct, portal vein, and inferior vena cava were selectively labeled with metallic nanoparticle (MCNP) dyes of different colors.

Hepatic artery: MCNP-Yellow was injected slowly via the left ventricle using a micro-syringe (25 μL, outer needle diameter 0.31 mm, Shanghai Bolige Industrial Co., Ltd.) until complete arterial filling or resistance occurred.

Bile duct: After blunt separation of the common bile duct and portal vein, MCNP-Green was retrogradely injected via the extrahepatic bile duct until surface bile ducts were fully filled or resistance was felt. The bile duct was ligated with 5-0 suture.

Portal vein: MCNP-Pink was injected via the portal vein using a micro-syringe until filling or resistance, followed by ligation.

Inferior vena cava: MCNP-Black was retrogradely injected via the inferior vena cava, followed by ligation.

Earlier version: Step 6: The liver was harvested and fixed at 4°C for 48 hours before proceeding with tissue clearing procedures.

Revised version: Step 6: The liver was harvested and fixed in 4% paraformaldehyde at 4°C for 48 hours, after which tissue clearing procedures were performed.

Earlier version: Step 7: Labeled liver tissues were sectioned to desired thickness and imaged using a high- magnification deep-focus microscope (Keyence), confocal laser scanning microscope

Revised version: Step 7: Labeled liver tissues were sectioned to desired thickness and imaged using a high-magnification deep-focus microscope (Keyence), confocal laser scanning microscope (Zeiss LSM980), or THUNDER high-resolution imaging system (Leica).

Removed in the revised version: (Zeiss LSM980), or THUNDER high-resolution imaging system (Leica).

Bile Duct High-Pressure Dye Perfusion:

Follow steps 1–4 as above. Then, using a micro-syringe, retrograde catheterization was performed via the extrahepatic bile duct with secure positioning at the hepatic hilar bile duct. A large volume of MCNP-Green was rapidly injected until the entire bile duct system and peripheral branches were filled, evident from surface dye visualization. Dye volume was adjusted to liver size, typically 25–40 μL per mouse.

Traditional CUBIC Reagent Preparation and Tissue Clearing

Based on Hiroki R. Ueda’s protocol [15]:

Step 1: Prepare ScaleCUBIC-1 (25% CUBIC-I): 25 wt% urea, 25 wt% N,N,N′,N′-tetrakis(2-hydroxypropyl)ethylenediamine, and 15 wt% Triton X-100. Prepare ScaleCUBIC-2 (25% CUBIC-II): 50 wt% sucrose, 25 wt% urea, 10 wt% triethanolamine, and 0.1% (v/v) Triton X-100.

Step 2: Fixed liver lobes were immersed in 20 g of CUBIC-I at 37°C with gentle shaking for 3 days, then replaced with fresh CUBIC-I for another 2 days. Samples were rinsed with PBS and immersed in CUBIC-II for 3 days.

Liver-CUBIC (40%CUBIC+H2O2) Reagent Preparation and Tissue Clearing

Step 1: Prepare 40%CUBIC-I: 36 wt% urea, 25 wt% N,N,N′,N′-tetrakis(2-hydroxypropyl)ethylenediamine, and 15 wt% Triton X-100. Prepare 40%CUBIC-II: 50 wt% sucrose, 35 wt% urea, 10 wt% triethanolamine, and 0.1% Triton X-100. Bleaching solution: 4.5% hydrogen peroxide and 24 mM sodium hydroxide.

Earlier version: Step 2: Whole liver lobes were immersed in IRISKit® HyperView Quench Buffer (#MH010301) at room temperature for 24 h with solution refreshed every 12 h. Samples were then washed with PBS and incubated in 20 g of 40%CUBIC-I at 37°C for 1 day, washed again, and finally incubated in 40%CUBIC-II at 37°C for 1 day.

Revised version: Step 2: We subjected all hepatic lobes of the whole liver (including the left, right, caudate, and quadrate lobes) to the Liver-CUBIC aqueous clearing protocol to ensure uniform visualization of MCNP fluorescence and immunolabeling signals throughout the three-dimensional reconstruction of the entire liver. The liver lobes were immersed in IRISKit® HyperView Quench Buffer (Cat. #MH010301) at room temperature for 24 h, with the solution refreshed every 12 h. The samples were then washed with PBS and incubated in 20 g of 40%CUBIC-I solution at 37 °C for 1 day, washed again, and subsequently incubated in 40%CUBIC-II solution at 37 °C for another 1 day.

Transmission Quantification

Added in the revised version: Transmittance was measured using a microplate reader (BioTek, USA) on 1 mm-thick liver sections.

Earlier version: Transmittance was measured using a microplate reader (BioTek, USA) on 1 mm-thick liver sections. Step 1: 1 mm liver slices were cut into 5 mm × 5 mm squares and treated with various clearing protocols, then immersed in CUBIC-II. Five replicates per group were placed into 96-well plates.

Revised version: Step 1: 1 mm liver slices were cut into 5 mm × 5 mm squares and treated with various clearing protocols, then immersed in CUBIC-II. Five replicates per group were placed into 96-well plates.

Step 2: Optical transmittance was recorded from 400 nm to 900 nm, generating transmittance spectra for each sample.

Added in the revised version: Transmission Electron Microscopy (TEM) of MCNPs

Added in the revised version: Prior to imaging, MCNP dye solutions were centrifuged at 14,000 × g for 10 min at 4 °C to remove aggregates and impurities. The supernatant was diluted 50-fold, and 3–4 μL of the diluted sample was applied onto freshly glow-discharged Quantifoil R1.2/1.3 copper grids (Electron Microscopy Sciences, 300 mesh). After a 30 s incubation to allow particle adsorption, excess liquid was gently blotted with filter paper and the grid was air-dried at room temperature. Samples were negatively stained with 1% uranyl acetate for 30 s and air-dried again before imaging.

Added in the revised version: Negative-stain TEM images were acquired using a 120 kV JEOL JEM-1400 transmission electron microscope (JEOL Ltd., Japan) equipped with a CCD camera. Data collection was performed using standard imaging conditions.

3D DAB Staining and Imaging

Earlier version: Step 1: Liver tissues were fixed in 4% paraformaldehyde at 4°C for 24 h, sectioned to 200 μm using a vibratome (Leica VT1200, Germany), and washed three times in distilled water (10 min each).

Revised version: Step 1: Liver tissues were fixed in 4% paraformaldehyde at 4°C for 24 h, sectioned to 200 μm using a vibratome (Leica VT1200, Germany), and washed three times in distilled water (10 min each), DAB staining was performed using free-floating sections. The sections were processed while floating in solution throughout the entire staining procedure.

Earlier version: Step 2: Sections were delipidated in acetone:chloroform (1:1, v/v) for 3 h.

Revised version: Step 2: Sections were delipidated in acetone: chloroform (1:1, v/v) at RT for 3 h.

Step 3: Dehydrated in graded ethanol (75%, 85%, 95%, 100%, 100%, 95%, 85%, 75%), 30 min each, followed by three 10 min water washes.

Earlier version: Step 4: Quenching was performed with quench reagent (including H2O2, IRISKit HyperView Quench Buffer) for 20 min in a LUMINIRIS quencher, followed by three water washes. Sections were treated with 2% EDTA antigen retrieval buffer (in distilled water) at 37°C for 35 min, then washed with PBS + 0.3% Triton X-100 three times (10 min each).

Revised version: Step 4: Quenching was performed with quench reagent (including H2O2, IRISKit HyperView Quench Buffer) for 20 min in a LUMINIRIS quencher, followed by three water washes. Sections were treated with 2% EDTA antigen retrieval buffer (in distilled water) at 95°C for 35 min, then washed with PBS + 0.3% Triton X-100 three times (10 min each).

Step 5: Incubated with primary antibody (rabbit anti-CK19, 1:2000, Huaan Biotechnology ET1601-6) at 4°C for 24 h, followed by three PBS washes.

Step 6: Incubated with polymerized HRP-conjugated secondary antibody at 4°C for 24 h, followed by PBS + 0.3% Triton washes.

Earlier version: Step 7: Stained with DAB (Zhongshan Jinqiao ZLI-9017) for 10 min, followed by three PBS washes. Step 8: Sequentially cleared with 40%CUBIC-I at 37°C for 30 min, washed, and incubated in 40%CUBIC-II at 37°C for 30 min.

Revised version: Step 7: Stained with DAB (Zhongshan Jinqiao ZLI-9017) at RT for 10 min, followed by three PBS washes.

Added in the revised version: Step 8: Sequentially cleared with 40%CUBIC-I at 37°C for 30 min, washed, and incubated in 40%CUBIC-II at 37°C for 30 min.

Step 9: 3D imaging was performed using a high-depth focus microscope.

3D TSA Multiplex Immunofluorescence Staining

Earlier version: Performed using IRISKit® HyperView multiplex immunostaining kit (#MH010101). Steps 1–4: Same as 3D DAB protocol.

Revised version: Performed using IRISKit® HyperView multiplex immunostaining kit (#MH010101).

Added in the revised version: Steps 1–4: Same as 3D DAB protocol.

Step 5: Incubated with primary antibodies (anti-CK19, CD31, HNF4α, ZO-1, CD34, Sca-1, and TH) overnight at 4°C.

Step 6: Incubated with polymerized HRP-conjugated secondary antibody overnight at 4°C.

Step 7: Sequentially visualized each antibody with specific TSA fluorophores: Cyclic-480 (EX 465–495 nm, EM 512–558 nm), Cyclic-550 (EX 540–560 nm, EM 575–595 nm), and Cyclic-630 (EX 620–640 nm, EM >665 nm).

Step 8: After each round, antibodies were stripped using IRISKit® HyperView advanced antibody removal kit.

Step 9: Cleared tissue was imaged by confocal laser scanning microscopy (Zeiss LSM980).

Added in the revised version: Imaging Equipment and Parameter Settings

Added in the revised version: In this study, three types of microscopes were employed for imaging analysis. The specific models and parameter settings are described as follows:

Added in the revised version: Ultra-Depth Microscope (VHX-6000, Keyence)

Added in the revised version: This microscope was equipped with a VH-Z100R objective lens, offering a magnification range of 100×–1000× and a typical resolution of 1 μm. The illumination mode was set to coaxial reflected light, and the platform’s transmitted illumination was turned “ON” to ensure uniform brightness across the tissue samples, thereby facilitating clear visualization of structural details.

Added in the revised version: Zeiss Confocal Microscope (LSM980)

Added in the revised version: Depending on imaging requirements, either a 20× or 40× objective was used, with the image size set to 1024 × 1024 pixels. Three fluorescence detection channels were configured as follows:

Added in the revised version: Channel 1: 639 nm laser excitation (λ_ex = 650 nm; λ_em = 517 nm; detection range = 673–758 nm), corresponding to Cy5-T1 (red fluorescence signal);

Added in the revised version: Channel 2: 561 nm laser excitation (λ_ex = 548 nm; λ_em = 561 nm; detection range = 547–637 nm), corresponding to Cy3-T2 (orange fluorescence signal);

Added in the revised version: Channel 3: 488 nm laser excitation (λ_ex = 493 nm; λ_em = 517 nm; detection range = 490–529 nm), corresponding to AF488-T3 (green fluorescence signal).

Added in the revised version: Leica THUNDER Imager 3D Tissue Microscope

Added in the revised version: This microscope was configured with two fluorescence detection channels:

Added in the revised version: Channel 1: FITC channel (λ_ex = 488 nm; λ_em ≈ 520 nm);

Added in the revised version: Channel 2: Orange-red channel (λ_ex = 561 nm; λ_em ≈ 640 nm).

Added in the revised version: To ensure signal specificity, a dedicated set of filters matched to the spectral characteristics of both channels was applied, effectively minimizing crosstalk between different fluorescent signals.

Added in the revised version: 3D Image Quantification

Added in the revised version: ImageJ software was used to analyze 3D stained images (Figs. 5F, 6E), and the ultra-depth-of-field 3D analysis module was used to analyze 3D DAB images (Fig. 5D). The specific steps are as follows:

Added in the revised version: Figure 5D: DAB-stained 3D images from the control group and the CCl4 6-week (CCl4-6W) group were analyzed. For each group, 20 terminal bile duct branch nodes were randomly selected, and the actual path distance along the branch to the nearest portal vein surface was measured. All measurements were plotted as scatter plots to reflect the spatial extension of bile ducts relative to the portal vein under different conditions.

Added in the revised version: Figure 5F: TSA 3D multiplex-stained images from the control group, CCl4 3-week (CCl4-3W), and CCl4 6-week (CCl4-6W) groups were analyzed. For each group, 5 terminal bile duct branch nodes were randomly selected, and the actual path distance along the branch to the nearest portal vein surface was measured. Measurements were plotted as scatter plots to illustrate bile duct spatial extension.

Added in the revised version: Figure 6E: TSA 3D multiplex-stained images from the control, CCl4-3W, and CCl4-6W groups were analyzed. For each group, 5 terminal nerve branch nodes were randomly selected, and the actual path distance along the branch to the nearest portal vein surface was measured. Scatter plots were generated to depict the spatial distribution of nerves under different treatment conditions.

Single-Cell Data Mining

Earlier version: We re-analyzed published single-cell transcriptomic datasets of liver endothelial cells (ECs) sorted by FACS from three normal mouse livers [22].

Revised version: We re-analyzed published single-cell transcriptomic datasets of liver endothelial cells (ECs) sorted by FACS from three normal mouse livers (Su et al., 2021). In this study, we downloaded single-cell RNA sequencing (scRNA-seq) data published by Su et al. from a public database and preprocessed the data according to the filtering criteria reported in the original study, retaining only GFP-positive cells.

Added in the revised version: Subsequently, standard analysis procedures were performed using Seurat (v5.0.3), including quality control, normalization, feature selection, dimensionality reduction, and clustering. Cell populations were annotated based on marker genes reported in the literature, and UMAP was used to visualize cell distribution. We further quantified the abundance of different endothelial cell (EC) subpopulations. After extracting the Seurat object for ECs, clustering results were corrected and re-clustered. Differential expression analysis was performed using the FindMarkers function. During this process, we selected the intersection of highly variable genes and differentially expressed genes in ECs as a candidate gene set and generated a heatmap showing the top 20 differentially expressed genes from portal vein ECs to periportal LSECs. Within this intersecting gene set, Cd34 and Sca-1 were identified. We then extracted the Cd34+Sca-1+ double-positive subpopulation from both portal vein and periportal vein clusters and continued analysis and differential gene screening according to the Seurat workflow.

Added in the revised version: Finally, the differentially expressed genes (p_val < 0.05 and avg_log2FC > 0.3) were subjected to Gene Ontology (GO) enrichment analysis with ont = “BP”. Enrichment bar plots and volcano plots were generated using ggplot2 (v3.5.1) to illustrate key differentially expressed genes and their functional enrichment.

Added in the revised version: To validate these results, we downloaded data from another public dataset containing cell populations similar to the PLC in our study. After extracting the target population (e2), cells were divided into Cd34-and Sca-1-positive and -negative groups, and the previous analyses—including differential expression and enrichment analyses—were repeated for validation(Pietilä et al., 2025).

Statistics and Reproducibility

Statistical analyses were performed using Prism 8.0 (GraphPad Software, Inc.). p-values are indicated in figure legends. Data are presented as mean ± SD unless otherwise stated. Two-group comparisons were made using Student’s t-test. One-way ANOVA followed by Tukey’s multiple comparisons test was used for multi-group analysis. Statistical significance was defined as ***p < 0.001, **p < 0.01, *p < 0.05.